Fractal cycles of sleep: a new aperiodic activity-based definition of sleep cycles

Author response:

The following is the authors’ response to the previous reviews.

Reviewer 1:

Weaknesses:

The match between fractal and classical cycles is not one-to-one. For example, the fractal method identifies a correlation between age and cycle duration in adults that is not apparent with the classical method. This raises the question as to whether differences are due to one method being more reliable than another or whether they are also identifying different underlying biological differences. It is not clear for example whether the agreement between the two methods is better or worse than between two human scorers, which generally serve as a gold standard to validate novel methods. The authors provide some insight into differences between the methods that could account for differences in results. However, given that the fractal method is automatic it would be important to clearly identify criteria for recordings in which it will produce similar results to the classical method.

We thank the reviewer for the insightful suggestions. In the revised Manuscript, we have added a number of additional analyses that provide a quantitative comparison between the classical and fractal cycle approaches aiming to identify the source of the discrepancies between classical and fractal cycle durations. Likewise, we assessed the intra-fractal and intra-classical method reliability.

Reviewer 2:

One weakness of the study, from my perspective, was that the IRASA fits to the data (e.g. the PSD, such as in Figure 1B), were not illustrated. One cannot get a sense of whether or not the algorithm is based entirely on the fractal component or whether the oscillatory component of the PSD also influences the slope calculations. This should be better illustrated, but I assume the fits are quite good.

Thank you for this suggestion. In the revised Manuscript, we have added a new figure (Fig.S1 E, Supplementary Material 2), illustrating the goodness of fit of the data as assessed by the IRASA method.

The cycles detected using IRASA are called fractal cycles. I appreciate the use of a simple term for this, but I am also concerned whether it could be potentially misleading? The term suggests there is something fractal about the cycle, whereas it's really just that the fractal component of the PSD is used to detect the cycle. A more appropriate term could be "fractal-detected cycles" or "fractal-based cycle" perhaps?

We agree that these cycles are not fractal per se. In the Introduction, when we mention them for the first time, we name them “fractal activity-based cycles of sleep” and immediately after that add “or fractal cycles for short”. In the revised version, we renewed this abbreviation with each new major section and in Abstract. Nevertheless, given that the term “fractal cycles” is used 88 times, after those “reminders”, we used the short name again to facilitate readability. We hope that this will highlight that the cycles are not fractal per se and thus reduce the possible confusion while keeping the manuscript short.

The study performs various comparisons of the durations of sleep cycles evaluated by the IRASA-based algorithm vs. conventional sleep scoring. One concern I had was that it appears cycles were simply identified by their order (first, second, etc.) but were not otherwise matched. This is problematic because, as evident from examples such as Figure 3B, sometimes one cycle conventionally scored is matched onto two fractal-based cycles. In the case of the Figure 3B example, it would be more appropriate to compare the duration of conventional cycle 5 vs. fractal cycle 7, rather than 5 vs. 5, as it appears is currently being performed.

In cases where the number of fractal cycles differed from the number of classical cycles (from 34 to 55% in different datasets as in the case of Fig.3B), we did not perform one-to-one matching of cycles. Instead, we averaged the duration of the fractal and classical cycles over each participant and only then correlated between them (Fig.2C). For a subset of the participants (45 – 66% of the participants in different datasets) with a one-to-one match between the fractal and classical cycles, we performed an additional correlation without averaging, i.e., we correlated the durations of individual fractal and classical cycles (Fig.4S of Supplementary Material 2). This is stated in the Methods, section Statistical analysis, paragraph 2.

There are a few statements in the discussion that I felt were either not well-supported. L629: about the "little biological foundation" of categorical definitions, e.g. for REM sleep or wake? I cannot agree with this statement as written. Also about "the gradual nature of typical biological processes". Surely the action potential is not gradual and there are many other examples of all-or-none biological events.

In the revised Manuscript, we have removed these statements from both Introduction and Discussion.

The authors appear to acknowledge a key point, which is that their methods do not discriminate between awake and REM periods. Thus their algorithm essentially detected cycles of slow-wave sleep alternating with wake/REM. Judging by the examples provided this appears to account for both the correspondence between fractal-based and conventional cycles, as well as their disagreements during the early part of the sleep cycle. While this point is acknowledged in the discussion section around L686. I am surprised that the authors then argue against this correspondence on L695. I did not find the "not-a-number" controls to be convincing. No examples were provided of such cycles, and it's hard to understand how positive z-values of the slopes are possible without the presence of some wake unless N1 stages are sufficient to provide a detected cycle (in which case, then the argument still holds except that its alterations between slow-wave sleep and N1 that could be what drives the detection).

In the revised Manuscript, we have removed the “NaN analysis” from both Results and Discussion. We have replaced it with the correlation between the difference between the durations of the classical and fractal cycles and proportion of wake after sleep onset. The finding is as follows:

“A larger difference between the durations of the classical and fractal cycles was associated with a higher proportion of wake after sleep onset in 3/5 datasets as well as in the merged dataset (Supplementary Material 2, Table S10).” Results, section “Fractal cycles and wake after sleep onset”, last two sentences. This is also discussed in Discussion, section “Fractal cycles and age”, paragraph 1, last sentence.

To me, it seems important to make clear whether the paper is proposing a different definition of cycles that could be easily detected without considering fractals or spectral slopes, but simply adjusting what one calls the onset/offset of a cycle, or whether there is something fundamentally important about measuring the PSD slope. The paper seems to be suggesting the latter but my sense from the results is that it's rather the former.

Thank you for this important comment. Overall, our paper suggests that the fractal approach might reflect the cycling nature of sleep in a more precise and sensitive way than classical hypnograms. Importantly, neither fractal nor classical methods can shed light on the mechanism underlying sleep cycle generation due to their correlational approach. Despite this, the advantages of fractal over classical methods mentioned in our Manuscript are as follows:

(1) Fractal cycles are based on a real-valued metric with known neurophysiological functional significance, which introduces a biological foundation and a more gradual impression of nocturnal changes compared to the abrupt changes that are inherent to hypnograms that use a rather arbitrary assigned categorical value (e.g., wake=0, REM=-1, N1=-2, N2=-3 and SWS=-4, Fig.2 A).

(2) Fractal cycle computation is automatic and thus objective, whereas classical sleep cycle detection is usually based on the visual inspection of hypnograms, which is time-consuming, subjective and error-prone. Few automatic algorithms are available for sleep cycle detection, which only moderately correlated with classical cycles detected by human raters (r’s = 0.3 – 0.7 in different datasets here).

(3) Defining the precise end of a classical sleep cycle with skipped REM sleep that is common in children, adolescents and young adults using a hypnogram is often difficult and arbitrary. The fractal cycle algorithm could detect such cycles in 93% of cases while the hypnogram-based agreement on the presence/absence of skipped cycles between two independent human raters was 61% only; thus, 32% lower.

(4) The fractal analysis showed a stronger effect size, higher F-value and R-squared than the classical analysis for the cycle duration comparison in children and adolescents vs young adults. The first and second fractal cycles were significantly shorter in the pediatric compared to the adult group, whereas the classical approach could not detect this difference.

(5) Fractal – but not classical – cycle durations correlated with the age of adult participants.

These bullets are now summarized in Table 5 that has been added to the Discussion of the revised manuscript.

Reviewer #1 (Recommendations for the authors):

The authors have added a lot of quantifications to provide a more complete comparison of classical and fractal cycles that address the points I raised.

Regarding, the question of skipped REM cycles: I am not sure the comparison of skipped cycle accuracies between fractal and manual methods makes sense. To make a fair comparison fractal and 2nd scorer classifications should be compared to the same baseline dataset which doesn't seem to be the case since the number of skipped cycles is not the same. Moreover, it's not indicated whether the fractal method identifies any false positive skipped cycles.

Thank you for this comment. In the revised Manuscript, we have reported the number of false positive skipped cycles identified by the fractal algorithm. Likewise, we have added the comparison between the fractal algorithm and the second scorer detection of cycles with skipped REM sleep (Results, the section “Skipped cycles”, last paragraph). The text has been revised as follows:

“Visual inspection of the hypnograms from Datasets 1 – 6 was performed by two independent researchers. Scorer 1 and Scorer 2 detected that out of 226 first sleep cycles 58 (26%) and 64 (28%), respectively, lacked REM episodes. The agreement on the presence of skipped cycles between two human raters equaled 91% (58 cycles detected by both raters out of 64 cycles detected by either one or two scorers). The fractal cycle algorithm detected skipped cycles in 57 out of 58 (98%) cases detected by Scorer 1 with one false positive (which, however, was tagged as a skipped cycle by Scorer2), and in 58 out of 64 (91%) cases detected by Scorer 2 with no false positives.”

Minor points

I suggest reporting the values of inter-method / inter-scorer correlations with the classical method in the main text since otherwise interpreting the value for fractal vs classical is impossible.

Thank you for this comment. In the revised Manuscript, we have moved this section to the main text (Table 3).

Table 5 + text of discussion: cycle identification based on hypnograms is claimed to be. "based on arbitrary assigned categorical values" the categories are not arbitrary since they correspond to well-validate sleep states, only the number associated it and this does not seem to be very important since it's only for visualization purposes.

Thank you for this comment. In the revised Manuscript, we have removed the phrase “arbitrary assigned“.

Changes from the earlier version to the revised version

Abstract

Abstract

Nocturnal human sleep consists of 4 – 6 ninety-minute cycles defined as episodes of non-rapid eye movement (non-REM) sleep followed by an episode of REM sleep. While sleep cycles are considered fundamental components of sleep, their functional significance largely remains unclear. One of the reasons for a lack of research progress in this field is the absence of a data-driven definition of sleep cycles. Here, we proposed to base such a definition on fractal (aperiodic) neural activity, a well-established marker of arousal and sleep stages.

We explored temporal dynamics of fractal activity during nocturnal sleep using electroencephalography. Based on the observed pattern of fractal fluctuations, we introduced a new concept of fractal activity-based cycles of sleep or “fractal cycles” for short, defined as a time interval during which fractal activity descends from its local maximum to its local minimum and then leads back to the next local maximum. Next, we assessed correlations between fractal and classical (i.e., non-REM – REM) sleep cycle durations. We also studied cycles with skipped REM sleep, i.e., the cycles where the REM phase is expected to appear except that it does not, being replaced by lightening of sleep.

Regarding the sample, we examined fractal cycles in healthy adults (age range: 18 – 75 years, n = 205) as well as in children and adolescents (range: 8 – 17 years, n = 21), the group characterized by deeper sleep and a higher frequency of cycles with skipped REM sleep. Further, we studied fractal cycles in major depressive disorder (n = 111), the condition characterized by altered REM sleep (in addition to its clinical symptoms).

Earlier version: We found that fractal and classical cycle durations (89 ± 34 min vs 90 ± 25 min) correlated positively (r = 0.5, p < 0.001). Cycle-to-cycle overnight dynamics showed an inverted U-shape of both fractal and classical cycle durations and a gradual decrease in absolute amplitudes of the fractal descents and ascents from early to late cycles. In adults, the fractal cycle duration and participant’s age correlated negatively (r = -0.2, p = 0.006). Children and adolescents had shorter fractal cycles compared to young adults (76 ± 34 vs 94 ± 32 min, p < 0.001). The fractal cycle algorithm detected cycles with skipped REM sleep in 90/97 (95%) cases while the hypnogram-based agreement on the presence of skipped REM sleep between two human raters was 61% only; thus, 32% lower. Medicated patients with depression showed longer fractal cycles compared to their own unmedicated state (107 ± 51 min vs 92 ± 38 min, p < 0.001) and age-matched controls (104 ± 49 vs 88 ± 31 min, p < 0.001).

Revised version: We found that fractal and classical cycle durations (89 ± 34 min vs 90 ± 25 min) correlated positively (r = 0.5, p < 0.001). Cycle-to-cycle overnight dynamics showed an inverted U-shape of both fractal and classical cycle durations and a gradual decrease in absolute amplitudes of the fractal descents and ascents from early to late cycles. In adults, the fractal cycle duration and participant’s age correlated negatively (r = −0.2, p = 0.006). Children and adolescents had shorter fractal cycles compared to young adults (76 ± 34 vs 94 ± 32 min, p < 0.001). The fractal cycle algorithm detected cycles with skipped REM sleep in 91 – 98% of cases. Medicated patients with depression showed longer fractal cycles compared to their own unmedicated state (107 ± 51 min vs 92 ± 38 min, p < 0.001) and age-matched controls (104 ± 49 vs 88 ± 31 min, p < 0.001).

In conclusion, fractal cycles are an objective, quantifiable, continuous and biologically plausible way to display sleep neural activity and its cycles. They are useful in healthy adult and pediatric populations as well as in patients with major depressive disorder. Fractal cycles should be extensively studied to advance theoretical research on sleep structure.

Introduction

The cyclic nature of sleep has long been established with a classical sleep cycle defined as a time interval that consists of an episode of non-rapid eye movement (non-REM) sleep followed by an episode of REM sleep (Feinberg & Floid, 1979; Le Bon, 2020). Typically, nocturnal sleep consists of 4 – 6 such cycles, which last for about 90 minutes each. Every cycle is seen as a fundamental physiological unit of sleep central to its function (Feinberg, 1974) or a miniature representation of the sleep process (Le Bon, 2002).

Basic structural organization of normal sleep is rather conservative with some exceptions. Thus, occasionally, at the beginning of the night in healthy adolescents and young adults, there could occur cycles with skipped REM sleep, which are also called “skipped” cycles. In skipped cycles, a REM sleep episode is expected to appear except that it does not and only a “lightening” of sleep is observed presumably due to too high non-REM pressure (Le Bon, 2020). Likewise, some alterations of the sleep structure can be observed in sleep disorders, e.g., narcolepsy and insomnia (Scammell, 2015), and healthy aging (Carrier et al., 2011; Conte et al., 2014). In some neurological and psychiatric conditions, such as major depressive disorder (MDD), Parkinson’s and Alzheimer’s diseases, sleep architecture disturbances are further linked to the disease neuropathology (Courtet & Olié, 2012; Palagini et al., 2013; Pillai & Leverenz, 2017).

While the importance of sleep cycles is indisputable, their function as a unit is poorly understood and surprisingly under-explored, especially when compared to the extensive research on sleep stages (either non-REM or REM) or sleep microstructure (e.g., sleep spindles, slow waves, microarousals). One of the reasons for this striking absence of research progress might be the lack of proper quantifiable and reliable objective measure from which sleep cycles could be derived directly (Schneider et al., 2022).

Currently, sleep cycles are defined via a visual inspection of the hypnogram, the graph in which categorically separated sleep stages are plotted over time. Yet assigning a discrete category to each sleep stage is rather arbitrary as sleep stages are presumably continuous and thus do not occur as steep lines of a hypnogram. In addition, visual sleep stage scoring is very time-consuming, subjective and error-prone with a relatively low (~80%) inter-rater agreement. This results in a low accuracy regarding the sleep cycle definition.

Earlier version: We suggest that a data-driven approach based on a real-valued neurophysiological metric (as opposed to the categorical one) with a finer quantized scale could forward the understanding of sleep cycles considerably. Specifically, we propose that research on sleep cycles would benefit from recent advances in the field of fractal neural activity. In literature, fractal activity is also called aperiodic, non-oscillatory, 1/f or scale-free activity, being named after the self-similarity exhibited by patterns of sensor signals across various time scales. Fractal activity is a distinct type of brain dynamics, which is sometimes seen as a “background” state of the brain, from which linear, rhythmic (i.e., periodic, oscillatory) dynamics emerge to support active processing (Buzsaki, 2006; Freeman et al., 2006). Growing evidence confirms that fractal activity has a rich information content, which opens a window into diverse neural processes associated with sleep, cognitive tasks, age and disease (Voytek & Knight, 2015; Bódizs et al., 2021; 2023; Höhn et al., 2022).

Revised version: We suggest that a data-driven approach based on a real-valued neurophysiological metric (as opposed to the categorical one) with a finer quantized scale could forward the understanding of sleep cycles considerably. Specifically, we propose that research on sleep cycles would benefit from recent advances in the field of fractal neural activity. In literature, fractal activity is also called aperiodic, non-oscillatory, 1/f or scale-free activity, being named after the self-similarity exhibited by patterns of sensor signals across various time scales. Fractal activity is a distinct type of brain dynamics, which is sometimes seen as a “background” state of the brain, from which linear, rhythmic (i.e., periodic, oscillatory) dynamics emerge to support active processing (Buzsaki, 2006; Freeman et al., 2006). Growing evidence confirms that fractal activity has a rich information content, which opens a window into diverse neural processes associated with sleep, cognitive tasks, age and disease (Voytek & Knight, 2015; Bódizs et al., 2021; 2024; Höhn et al., 2022).

Fractal dynamics follow a power-law 1/f function, where power decreases with increasing frequency (He, 2014). The steepness of this decay is approximated by the spectral exponent, which is equivalent to the slope of the spectrum when plotted in the log-log space (He, 2014; Gerster et al., 2022). The fractal signal is not dominated by any specific frequency, rather it reflects the overall frequency composition within the time series (Horváth et al., 2022) such that steeper (more negative) slopes indicate that the spectral power is relatively stronger in slow frequencies and relatively weaker in faster ones (He, 2014).

In terms of mechanisms, it has been suggested that flatter high-band (30 – 50Hz) fractal slopes reflect a shift in the balance between excitatory and inhibitory neural currents in favour of excitation while steeper slopes reflect a shift towards inhibition (Gao et al., 2017). Given that the specific balance between excitation and inhibition defines a specific arousal state and the conscious experience of an organism (Nir & Tononi, 2010), the introduction of Gao’s model led to an increased interest in fractal activity. For example, it has been shown that high-band fractal slopes discriminate between wakefulness, non-REM and REM sleep stages as well as general anesthesia or unconsciousness (Gao et al., 2017; Colombo et al., 2019; Lendner et al., 2020; Höhn et al., 2022).

Of note, Gao’s model does not account for the lower part of the spectrum, which is also scale-free. An alternative model suggests that the broadband 1/f² activity reflects the tendency of the central nervous system to alternate between UP- (very rapid spiking) and DOWN- (disfacilitation, no activity) states (Milstein et al., 2009; Baranauskas et al., 2012). Empirical studies further showed that the broadband (2 – 48Hz) slope is an especially strong indicator of sleep stages and sleep intensity with low inter-subject variability and sensitivity to age-related differences (Miskovic et al., 2019; Schneider et al., 2022; Horváth et al., 2022). Taken together, this literature suggests that fractal slopes can serve as a marker of arousal, sleep stages and sleep intensity (Lendner et al., 2020; Schneider et al., 2022; Horváth et al., 2022). We expect that this line of inquiry could be extended to sleep cycles.

On a related note, the reciprocal interaction model of sleep cycles assumes that each sleep stage involves distinct activation patterns of inhibitory and excitatory neural networks (Pace-Schott & Hobson, 2002). This model explains alternations between non-REM and REM sleep stages by the interaction between aminergic and cholinergic neurons of the mesopontine junction (Pace-Schott & Hobson, 2002). Notably, during REM sleep, acetylcholine plays a major role in maintaining brain activation, which is expressed as EEG desynchronization, one of the main features of REM sleep (Nir & Tononi, 2010). This is of special importance in affective disorders since according to one of the pathophysiological explanations of depression, i.e., the cholinergic-adrenergic hypothesis, central cholinergic factors play a crucial role in the aetiology of affective disorders, with depression being a disease of cholinergic dominance (Janowsky et al., 1972). Many antidepressants (e.g., serotonin-norepinephrine reuptake inhibitors, selective serotonin reuptake inhibitors) suppress REM sleep and thus cause essential alterations in sleep architecture. Intriguingly, REM sleep suppression is related to the improvement of depression during pharmacological treatment with antidepressants enhancing monoaminergic neurotransmission (Vogel et al., 1990; Wichniak et al., 2013).

Based on this background, we propose that a fractal neural activity-based definition of sleep cycles has the potential to considerably advance our understanding of the cyclic nature of sleep, for example, by introducing graduality to the categorical concept of sleep stages. The current study analyzes the dynamics of nocturnal fluctuations in fractal activity using five independently collected polysomnographic datasets overall comprising 205 recordings from healthy adults. Based on the inspection of fractal activity across a night, we introduce a new concept of fractal activity-based cycles of sleep or “fractal cycles” for short. We describe differences and similarities between fractal cycles defined by our algorithm and classical (non-REM – REM) cycles defined by the hypnogram. We hypothesize that the timing and durations of the fractal cycles would closely correspond to those of classical cycles. We had no prior hypothesis regarding correspondence between the fractal cycles and classical cycles with skipped REM sleep, i.e., this analysis was exploratory.

Given the above-mentioned age-related changes in fractal activity (flatter slopes) and sleep structure (fewer and shorter classical cycles), we also study whether fractal cycle characteristics change with age. To this end we use 5 healthy adult datasets with the age range of 18 – 75 years (n = 205). Moreover, we add to our study a pediatric polysomnographic dataset (age range: 8 – 17 years, n = 21) to explore fractal cycles in childhood and adolescence, a life period accompanied by deepest sleep and massive brain reorganization (Kurth et al., 2012) as well as a higher frequency of cycles with skipped REM sleep (Jenni & Carskadon, 2004).

Finally, we test the clinical value of the fractal cycles by analyzing polysomnographic data in 111 patients with MDD, a condition characterized by disturbed sleep structure (besides its clinical symptoms, such as abnormalities of mood and affect). Specifically, we compare fractal cycles of sleep between medicated MDD patients (three MDD datasets, n = 111) and healthy age-matched controls (n = 111) as well as in the unmedicated and medicated states within the same MDD patients (one of the three MDD datasets, n = 38). We hypothesize that the fractal cycle approach would be more sensitive in detecting differences between typical and atypical sleep architecture compared to the conventional classical cycles.

Methods

Healthy participants

We retrospectively analyzed polysomnographic recordings from the following studies (Table 1):

Earlier version: Datasets 1 – 3: 40, 40 and 33 healthy controls from three independent sleep studies in MDD conducted at the Max Planck Institute of Psychiatry, Germany. These datasets are described in Rosenblum et al. (2023 a) and Bovy et al. (2022). In addition, these participants are used as controls in MDD datasets A – C described below.

Revised version: Datasets 1 – 3: 40, 40 and 33 healthy controls from three independent sleep studies in MDD conducted at the Max Planck Institute of Psychiatry, Germany. These datasets are described in Rosenblum et al. (2023a) and Bovy et al. (2022). In addition, these participants are used as controls in MDD datasets A – C described below.

Dataset 4: 36 healthy participants from a home-based sleep study exploring simultaneous polysomnographic and EEG wearables conducted at the Donders Institute for Brain, Cognition and Behavior, the Netherlands (Described as Dataset 2 in Jafarzadeh Esfahani et al., 2023). The signal was recorded at participants’ homes over three nights with a gap of a week between each recording. For consistency with other datasets (i.e., to end up with a comparable number of cycles provided by each participant), we used polysomnography (and not EEG recorded by wearables) from the first night only since it had the largest sample size (i.e., 5 subjects dropped out from the study after the first polysomnographic recording).

Earlier version: Dataset 5: 68 healthy controls from previous endocrinological studies conducted at the Max Planck Institute of Psychiatry, Germany, using only nights with no pharmacological or endocrine intervention. 60/68 participants are described in Rosenblum et al. (2024 a).

Revised version: Dataset 5: 68 healthy controls from previous endocrinological studies conducted at the Max Planck Institute of Psychiatry, Germany, using only nights with no pharmacological or endocrine intervention. 60/68 participants are described in Rosenblum et al. (2024a).

Dataset 6: 21 healthy children and adolescents from previous studies (Furrer et al., 2019; Volk et al., 2019; Jaramillo et al., 2020) conducted at the University Children’s Hospital Zürich, Switzerland. For the control group to this dataset, we selected all healthy adults from Datasets 1 – 3, 5, 6 (n = 205) whose ages lay in the range of 23 – 25 years (the age when the brain maturation process is supposed to be finished (Giedd & Rapoport, 2010) and no age-related processes are expected to start). This resulted in 24 subjects with a mean age of 24.8 ± 0.9 years.

The studies were approved by the Ethics committee of the University of Munich (Datasets 1 – 3, 5), Radboud University (Dataset 4) and Canton of Zürich (Dataset 6). All participants (or participants’ parents for Dataset 6) gave written informed consent.

Patients with MDD

Earlier version: We retrospectively analyzed polysomnographic recordings from our previous studies (Bovy et al., 2022; Rosenblum et al., 2023 a, Tables 1 – 2):

Revised version: We retrospectively analyzed polysomnographic recordings from our previous studies (Bovy et al., 2022; Rosenblum et al., 2023a, Tables 1–2):

Dataset A: 40 long-term medicated MDD patients vs 40 age- and gender-matched healthy controls (Dataset 1 here).

Dataset B: 38 MDD patients in unmedicated and 7-day medicated states vs 40 healthy age and gender-matched controls (Dataset 2 here).

Dataset C: 33 MDD patients at 7-day and 28-day of medication treatment vs 33 healthy age and gender-matched controls (Dataset 3 here).

Earlier version: Demographic and sleep characteristics of the patients, medication treatment and polysomnographic devices are described in our previous works (Bovy et al., 2022; Rosenblum et al., 2023 a). Here, Table S5 (Supplementary Material 2) presents medication treatment. In Rosenblum et al. (2023 a), Datasets A, B and C are referred to as the Replication Dataset 2, Main Dataset and Replication Dataset 1, respectively; in Bovy et al. (2022), the naming is the same as here. All studies were approved by the Ethics committee of the University of Munich. All participants gave written informed consent.

Revised version: Demographic and sleep characteristics of the patients, medication treatment and polysomnographic devices are described in our previous works (Bovy et al., 2022; Rosenblum et al., 2023a). Here, Table S5 (Supplementary Material) presents medication treatment. In Rosenblum et al. (2023a), Datasets A, B and C are referred to as the Replication Dataset 2, Main Dataset and Replication Dataset 1, respectively; in Bovy et al. (2022), the naming is the same as here. All studies were approved by the Ethics committee of the University of Munich. All participants gave written informed consent.

The first part of this study analyzes the data from healthy participants only and labels the datasets with the numbers 1 – 6. The second part of this study compares patients and controls and labels the analyzed datasets with the letters A – C. Notably, healthy participants used as controls in datasets A – C are the same subjects analyzed in Datasets 1 – 3.

Earlier version: In Supplementary Material 2, we report how many participants and for what reasons were excluded from the analysis. An example of one excluded participant is given in Fig.S6 C (S37). Likewise, we report pilot findings on fractal cycles in patients with psychophysiological insomnia, using the open access dataset from Rezaei et al., 2017 (Fig.S10, Supplementary Material 2).

Revised version: In Supplementary Material, we report how many participants and for what reasons were excluded from the analysis. An example of one excluded participant is given in Fig.S6 C (S37). Likewise, we report pilot findings on fractal cycles in patients with psychophysiological insomnia, using the open access dataset from Rezaei et al., 2017 (Fig.S10, Supplementary Material).

Polysomnography

Information about the studies and polysomnographic devices is reported in Table 1. The participants slept wearing a polysomnographic device in a sleep laboratory (Datasets 1 – 3, 5, 6) or in the home environment (Dataset 4). In datasets 1 – 3 and 5 all participants had an adaptation night before the examination night; adaptation night data was not available to be analyzed and reported here. In dataset 6, all participants had two recording nights: a baseline and an examination night with auditory stimulation. Here, only the baseline night was analyzed, which was either the first night (in 50% of cases) or the second night for a given participant.

Sleep stages were previously scored manually by independent experts according to the AASM standards (AASM, 2014). In the pediatric dataset, we used 20-s epochs, in the rest of the datasets, we used 30-s epochs. Epochs with EMG and EEG artifacts and channels with more than 20% artifacts during non-REM sleep were manually excluded by an experienced scorer before all automatic analyses.

Earlier version: We opted to analyze the F3 and F4 electrodes for maximal consistency between the studies as these leads were available in 6 out of 7 datasets. Another reason is that in our future studies, we plan to replicate this work using the data recorded with at-home wearable devices, which often have only frontal channels (e.g., F7 and F8). We report the topographical analysis over central, parietal and occipital electrodes (when available) in healthy and clinical datasets in Tables S1 and S6 of Supplementary Material 2, respectively, showing comparable results. In Table S1 of Supplementary Material 2, we also report correlations between fractal cycle durations defined using different channels.

Revised version: We opted to analyze the F3 and F4 electrodes for maximal consistency between the studies as these leads were available in 6 out of 7 datasets. Another reason is that in our future studies, we plan to replicate this work using the data recorded with at-home wearable devices, which often have only frontal channels (e.g., F7 and F8). We report the topographical analysis over central, parietal and occipital electrodes (when available) in healthy and clinical datasets in Tables S1 and S6 of Supplementary Material, respectively, showing comparable results. In Table S1 of Supplementary Material, we also report correlations between fractal cycle durations defined using different channels.

Fractal power component

Earlier version: Offline EEG data analyses were carried out with MATLAB (version R2021b, The MathWorks, Inc., Natick, MA), using the Fieldtrip toolbox and custom-made scripts. For each participant, we averaged the EEG signal over the F3 and F4 electrodes (or C3 and C4 – for Dataset 1 where the frontal channels were unavailable), calculated its spectral power for every 30 (adult datasets) or 20 (the pediatric dataset) seconds corresponding to the conventionally defined duration of sleep epochs and differentiated the total power to its fractal (i.e., aperiodic, 1/f, scale-free) and oscillatory components. Several methods to calculate fractal components exist. We opted to use the Irregularly Resampled Auto-Spectral Analysis (IRASA; Wen & Liu, 2016) tool embedded in the Fieldtrip toolbox (Oostenveld et al. 2011), one of the leading open-source EEG softwares, with the ft_freqanalysis function as described elsewhere (Rosenblum et al., 2023 a; 2023 b). A side note: slopes calculated with the IRASA strongly correlate (r = |0.9|) with those calculated using the “fitting oscillations and one over f” (FOOOF, See Supplementary Material in Schneider et al., 2022), another useful method used for aperiodic analysis (Donoghue et al., 2020). The fractal power component (shown in Fig.S1 A of Supplementary Material 2) was transformed to log-log coordinates and its slope was calculated to estimate the power-law exponent (the rate of spectral decay), using the function logfit (Lansey, 2020). The loglog data fit is shown in Fig.S1 E of Supplementary Material 2. The analysis flowchart is depicted in Fig.1 A; outputs of some of the analysis steps in an example individual are shown in Fig.1 B.

Revised version: Offline EEG data analyses were carried out with MATLAB (version R2021b, The MathWorks, Inc., Natick, MA), using the Fieldtrip toolbox and custom-made scripts. For each participant, we averaged the EEG signal over the F3 and F4 electrodes (or C3 and C4 – for Dataset 1 where the frontal channels were unavailable), calculated its spectral power for every 30 (adult datasets) or 20 (the pediatric dataset) seconds corresponding to the conventionally defined duration of sleep epochs and differentiated the total power to its fractal (i.e., aperiodic, 1/f, scale-free) and oscillatory components. Several methods to calculate fractal components exist. We opted to use the Irregularly Resampled Auto-Spectral Analysis (IRASA; Wen & Liu, 2016) tool embedded in the Fieldtrip toolbox (Oostenveld et al. 2011), one of the leading open-source EEG softwares, with the ft_freqanalysis function as described elsewhere (Rosenblum et al., 2023a; 2023b). A side note: slopes calculated with the IRASA strongly correlate (r = |0.9|) with those calculated using the “fitting oscillations and one over f” (FOOOF, See Supplementary Material in Schneider et al., 2022), another useful method used for aperiodic analysis (Donoghue et al., 2020). The fractal power component (shown in Fig.S1 A of Supplementary Material) was transformed to log-log coordinates and its slope was calculated to estimate the power-law exponent (the rate of spectral decay), using the function logfit (Lansey, 2020). The loglog data fit is shown in Fig.S1 E of Supplementary Material. The analysis flowchart is depicted in Fig.1 A; outputs of some of the analysis steps in an example individual are shown in Fig.1 B.

Analysis.

A. Analysis flowchart. IRASA – Irregularly Resampled Auto-Spectral Analysis, sgolayfilt – Savitzky-Golay filter. B. Outputs of some of the analysis steps in an example healthy 26-year-old individual. From top to bottom: time-frequency representation of the total spectral power, raw and smoothed time series of the fractal slopes and hypnogram. Frontal spectral power and its slopes were calculated in the 0.3 – 30 Hz range for each 30 seconds of sleep.

As opposed to the oscillatory component, the fractal component is usually treated as a unity and, therefore, is filtered in the broadband frequency range (Donoghue et al., 2020; Bódizs et al., 2021; Gerster et al., 2022). Nevertheless, different studies defined (slightly) differing bands, e.g., 30 – 50Hz (Gao et al., 2017; Lendner et al., 2020), 3 – 55Hz (Waschke et al., 2021), 0.5 – 35Hz (Miskovic et al., 2019), 1 – 40Hz, 1 – 20Hz and 20 – 40Hz (Colombo et al., 2019), 1 – 45Hz (Helson et al., 2023), 0.5 – 40Hz (Vinding et al., 2023), 3 – 45Hz and 30 – 45Hz (Höhn et al., 2022) and 2 – 48Hz (Bódizs et al., 2021; Schneider et al., 2022).

Earlier version: Here, we used the 0.3 – 30Hz range as this is a typical sleep frequency band used in many areas of sleep research, showing good ability to differentiate between sleep stage as could be seen in Fig.S1 B (Supplementary Material 2), which replicates existing literature. Dataset 4 was analyzed in the 0.3 – 18Hz range since relatively low low-pass filtering was applied to it during the recording (see Table 1). In Table S2 of Supplementary Material 2, we also analyze the 1 – 30Hz band to control for a possible distortion (the so called “knees’’ of the spectrum) of the linear fit by excluding low frequencies with strong oscillatory activity (Gao et al., 2017; Bódizs et al., 2021). We find that the results are similar to those obtained for the 0.3 – 30Hz band reported in the Main text (probably thanks to the smoothening procedure we applied).

Revised version: Here, we used the 0.3 – 30Hz range as this is a typical sleep frequency band used in many areas of sleep research, showing good ability to differentiate between sleep stage as could be seen in Fig.S1 B (Supplementary Material), which replicates existing literature. Dataset 4 was analyzed in the 0.3 – 18Hz range since relatively low low-pass filtering was applied to it during the recording (see Table 1). In Table S2 of Supplementary Material, we also analyze the 1 – 30Hz band to control for a possible distortion (the so called “knees’’ of the spectrum) of the linear fit by excluding low frequencies with strong oscillatory activity (Gao et al., 2017; Bódizs et al., 2021). We find that the results are similar to those obtained for the 0.3 – 30Hz band reported in the Main text (probably thanks to the smoothening procedure we applied).

Earlier version: Finally, Fig.S1 D (Supplementary Material 2) shows aperiodic slopes in the 30 – 48Hz band averaged over sleep stages for Datasets 1 – 3 and 5. According to literature, REM sleep is expected to show the steepest (most negative) high-band slopes compared to all other sleep stages. However, we were able to replicate this finding in Datasets 1 and 5 only. Given poor differentiation between the stages in 2/4 datasets, this variable was not used in any further analyses.

Revised version: Finally, Fig.S1 D (Supplementary Material) shows aperiodic slopes in the 30 – 48Hz band averaged over sleep stages for Datasets 1 – 3 and 5. According to literature, REM sleep is expected to show the steepest (most negative) high-band slopes compared to all other sleep stages. However, we were able to replicate this finding in Datasets 1 and 5 only. Given poor differentiation between the stages in 2/4 datasets, this variable was not used in any further analyses.

Fractal activity-based cycles of sleep

Earlier version: Fractal activity-based cycles of sleep or “fractal cycles” for short were defined from fractal slope time series. For this, time series of the fractal slopes were z-normalized (raw values can be seen in Fig.S1 C, Supplementary Material 2) within a participant and smoothened with the Savitzky-Golay filter (Fig.1), the filter highly used in many fields of data processing. We used the Matlab’s function sgolayfilt(slope_time_series, order, frame_length) with the polynomial order of five and the frame length of 101. The peaks of the smoothed time series of the fractal slopes were defined with Matlab’s function findpeaks (slope_time_series, ’MinPeakDistance’, 40, ’MinPeakProminence’, 0.9) with the minimum peak distance of 20 minutes (i.e., forty 30-second epochs) and minimum peak prominence of |0.9| z (Fig.2 A – B). The amplitude of the descending and ascending phases of a cycle was defined to be > |0.9| z, meaning that there is a probability of p=0.8 that a given fractal slope lies below/above the standard normal distribution.

Revised version: Fractal activity-based cycles of sleep or “fractal cycles” for short were defined from fractal slope time series. For this, time series of the fractal slopes were z-normalized (raw values can be seen in Fig.S1 C, Supplementary Material) within a participant and smoothened with the Savitzky-Golay filter (Fig.1), the filter highly used in many fields of data processing. We used the Matlab’s function sgolayfilt(slope_time_series, order, frame_length) with the polynomial order of five and the frame length of 101. The peaks of the smoothed time series of the fractal slopes were defined with Matlab’s function findpeaks (slope_time_series, ‘MinPeakDistance’, 40, ‘MinPeakProminence’, 0.9) with the minimum peak distance of 20 minutes (i.e., forty 30-second epochs) and minimum peak prominence of |0.9| z (Fig.2 A – B). The amplitude of the descending and ascending phases of a cycle was defined to be > |0.9| z, meaning that there is a probability of p=0.8 that a given fractal slope lies below/above the standard normal distribution.

Earlier version: Fractal cycles in healthy adults. A – B. Individual fractal and classical sleep cycles.

Revised version: Fractal cycles in healthy adults.

Earlier version: Time series of smoothed z-normalized fractal slopes (bottom) and corresponding hypnograms (top) observed in two different participants. The duration of the fractal cycle is a time interval between two successive peaks (blue diamonds). A: S15 from Dataset 3 shows a one-to-one match between fractal cycles defined by the algorithm and classical (non-REM – REM) cycles defined by the hypnogram. B: In S22 from dataset 5, the second part of night has many wake epochs, some of them are identified by the algorithm as local peaks. This results in a higher number of fractal cycles as compared to the classical ones and a poor match between the fractal cycles No. 3 – 7 and classical cycles No. 2 – 5. The algorithm does not distinguish between the wake and REM-related fractal slopes and can define both as local peaks. Since the duration of the fractal cycles is defined as an interval of time between two adjacent peaks, more awakenings/arousals during sleep (usually associated with aging, Fig.S5 B) are expected to result in more peaks and, consequently, more fractal cycles, i.e., a shorter cycle duration. This is one of the possible explanations for the correlation between the fractal cycle duration and age (shown in Fig. S5 A). Time series of the fractal slopes and corresponding hypnograms for all participants are reported in Supplementary Material 1. SWS – slow-wave sleep, REM – rapid eye movement. C. Scatterplots: each dot represents the duration of the cycles averaged over one participant. The durations of the fractal and classical sleep cycles averaged over each participant correlate in all analyzed datasets, raw (non-ranked) values are shown, r – Spearman’s correlation coefficient. D. Cycle-to-cycle overnight dynamics show an inverted U shape of the duration of both fractal and classical cycles across a night and a gradual decrease in absolute amplitudes of the fractal descents and ascents from early to late cycles.

Revised version: A – B. Individual fractal and classical sleep cycles. Time series of smoothed z-normalized fractal slopes (bottom) and corresponding hypnograms (top) observed in two different participants. The duration of the fractal cycle is a time interval between two successive peaks (blue diamonds). A: S15 from Dataset 3 shows a one-to-one match between fractal cycles defined by the algorithm and classical (non-REM – REM) cycles defined by the hypnogram. B:In S22 from dataset 5, the second part of night has many wake epochs, some of them are identified by the algorithm as local peaks. This results in a higher number of fractal cycles as compared to the classical ones and a poor match between the fractal cycles No. 3 – 7 and classical cycles No. 2 – 5. The algorithm does not distinguish between the wake and REM-related fractal slopes and can define both as local peaks. Since the duration of the fractal cycles is defined as an interval of time between two adjacent peaks, more awakenings/arousals during sleep (usually associated with aging, Fig.S5 B) are expected to result in more peaks and, consequently, more fractal cycles, i.e., a shorter cycle duration. This is one of the possible explanations for the correlation between the fractal cycle duration and age (shown in Fig. S5 A). Time series of the fractal slopes and corresponding hypnograms for all participants are reportedshwon in Supplementary PowerPoint File shared on https://osf.io/gxzyd. SWS – slow-wave sleep, REM – rapid eye movement. C. Scatterplots: each dot represents the duration of the cycles averaged over one participant. The durations of the fractal and classical sleep cycles averaged over each participant correlate in all analyzed datasets, raw (non-ranked) values are shown, r – Spearman’s correlation coefficient. D. Cycle-to-cycle overnight dynamics show an inverted U shape of the duration of both fractal and classical cycles across a night and a gradual decrease in absolute amplitudes of the fractal descents and ascents from early to late cycles.

Of note, we had no solid a priori theoretical indication for choosing either of the function settings mentioned above. All settings were chosen a posteriori following an exploratory visual inspection of the normalized data from one dataset (Dataset 5), which therefore can be transferred to other datasets. That is, in datasets 1 – 4 and 6, the settings of the sgolayfilt and findpeaks functions were defined a priori based on the results obtained while inspecting Dataset 5.

Earlier version: In Table S7 of Supplementary Material 2, we compare results obtained while using different thresholds of the abovementioned parameters; namely, longer and shorter smoothing windows and higher and lower minimum peak prominence.

Revised version: In Table S7 of Supplementary Material, we compare results obtained while using different thresholds of the abovementioned parameters; namely, longer and shorter smoothing windows and higher and lower minimum peak prominence.

Classical sleep cycles

Earlier version: Classical sleep cycles were defined manually via the visual inspection of the hypnograms by two independent scorers according to the criteria originally proposed by Feinberg and Floyd (1979) with some adaptations as follows. A cycle typically starts with N1, N2 or sometimes wake and is followed by N2 or N2 and slow-wave sleep (SWS) > 20 minutes in duration, which can include wake. The cycle ends with the end of the REM period, which can include wake or short segments of non-REM sleep. No minimum REM duration criterion was applied (Tarokh et al., 2012). In some cases (described below), the cycle end was defined at a non-REM sleep stage or wake. Two examples of hypnograms with marked classical sleep cycles are shown in Fig.2 A – B. Four more examples are presented in Fig.S2 (Supplementary Material 2).

Revised version: Classical sleep cycles were defined manually via the visual inspection of the hypnograms by two independent scorers according to the criteria originally proposed by Feinberg and Floyd (1979) with some adaptations as follows. A cycle typically starts with N1, N2 or sometimes wake and is followed by N2 or N2 and slow-wave sleep (SWS) > 20 minutes in duration, which can include wake. The cycle ends with the end of the REM period, which can include wake or short segments of non-REM sleep. No minimum REM duration criterion was applied (Tarokh et al., 2012). In some cases (described below), the cycle end was defined at a non-REM sleep stage or wake. Two examples of hypnograms with marked classical sleep cycles are shown in Fig.2 A – B. Four more examples are presented in Fig.S2 (Supplementary Material).

The last incomplete (not terminated by the REM sleep phase) cycle at the end of the night was included in the analysis if its duration was > 50 minutes. The last incomplete cycles < 50 minutes were removed (nevertheless, they are shown in figures when present).

Earlier version: In Supplementary Excel File, we report classical cycle durations for each participant as scored by two human raters and the automatic algorithm (Blume & Cajochen, 2021). In Table S8 of Supplementary Material 2, we report the inter-rater agreement in number and durations of classical cycles.

Revised version: In Supplementary Excel File shared on https://osf.io/gxzyd, we report classical cycle durations for each participant as scored by two human raters and the automatic algorithm (Blume & Cajochen, 2021). In Table S8 of Supplementary Material, we report the inter-rater agreement in number and durations of classical cycles.

Skipped cycles

Earlier version: Given the absence of strict and broadly accepted rules for cycles with skipped REM sleep definition in literature, here, we tagged a cycle as “skipped” based on the visual inspection of the hypnogram combined with the criteria proposed by Jenni and Carskadon (2004) and Tarokh et al. (2012). Specifically, we subdivided a long cycle > 110 minutes into two when: 1) there was a “lightening of sleep” (i.e., the presence of wake, N1 and N2) in the middle of the long cycle, when a REM sleep episode was anticipated, 2) a continuous episode of N1, N2, wake or movement time lasting at least 12 minutes was preceded and followed by slow-wave sleep (Jenni & Carskadon, 2004); 3) two clear episodes of slow-wave sleep were separated by lighter non-REM stages (which might include wake) (Campbell et al., 2011; Tarokh et al., 2012). Long cycles containing skipped cycles were divided into cycles at time of sleep lightening. Examples of hypnograms with skipped sleep are shown in Fig.S6 and Fig.S9 (Supplementary Material 2). For each dataset, we checked whether the classical cycles with skipped REM sleep had been detected by the fractal cycle algorithm.

Revised version: Given the absence of strict and broadly accepted rules for cycles with skipped REM sleep definition in literature, here, we tagged a cycle as “skipped” based on the visual inspection of the hypnogram combined with the criteria proposed by Jenni and Carskadon (2004) and Tarokh et al. (2012). Specifically, we subdivided a long cycle > 110 minutes into two when: 1) there was a “lightening of sleep” (i.e., the presence of wake, N1 and N2) in the middle of the long cycle, when a REM sleep episode was anticipated, 2) a continuous episode of N1, N2, wake or movement time lasting at least 12 minutes was preceded and followed by slow-wave sleep (Jenni & Carskadon, 2004); 3) two clear episodes of slow-wave sleep were separated by lighter non-REM stages (which might include wake) (Campbell et al., 2011; Tarokh et al., 2012). Long cycles containing skipped cycles were divided into cycles at time of sleep lightening. Examples of hypnograms with skipped sleep are shown in Fig.S6 and Fig.S9 (Supplementary Material). For each dataset, we checked whether the classical cycles with skipped REM sleep had been detected by the fractal cycle algorithm.

Earlier version: In Supplementary Excel File, we report which classical cycles were tagged as “skipped” by two human raters. In Supplementary Material 2, we report the inter-rater agreement in number of cycles with skipped REM sleep (Table S9). In Supplementary Material 1, hypnograms of all healthy adult participants are presented next to fractal cycles with skipped cycles marked individually as assessed by rater 1.

Revised version: In Supplementary Excel File shared on https://osf.io/gxzyd, we report which classical cycles were tagged as “skipped” by two human raters. In Supplementary Material, we report the inter-rater agreement in number of cycles with skipped REM sleep (Table S9). In Supplementary PowerPoint File shared on https://osf.io/gxzyd, hypnograms of all healthy adult participants are presented next to fractal cycles with skipped cycles marked individually as assessed by rater 1.

Statistical analysis

The assumption that durations of the fractal and classical cycles come from a standard normal distribution was tested using the one-sample Kolmogorov-Smirnov test. The result suggested that this assumption should be rejected (p<0.05); therefore, non-parametric tests were used for all further analyses.

We correlated fractal and classical cycle durations using Spearman’s correlations in each dataset separately as well as in all datasets pooled. Given that in some participants (from 34 to 55% in different datasets), the number of the fractal cycles (mean 4.6 ± 1.0 cycles per participant) was not equal to the number of the classical cycles (mean 4.7 ± 0.9 cycles per participant), prior to the correlation analysis, we averaged the duration of the fractal and classical cycles over each participant. For a subset of the participants (45 – 66% of the participants in different datasets) with a one-to-one match between the fractal and classical cycles, we performed an additional correlation without averaging, i.e., we correlated the durations of individual fractal and classical cycles.

Earlier version: To identify sources of fractal and classical cycle mismatch, in Supplementary Material 2, Table S10, we further correlated between the difference in classical vs fractal sleep cycle durations on the one side and either the amplitude of fractal descend/ascend (to reflect fractal cycle depth), duration of cycles with skipped REM sleep, duration of wake after sleep onset or the REM episode length of a given cycle (to reflect peak flatness) on the other side.

Revised version: To identify sources of fractal and classical cycle mismatch, we further correlated between the difference in classical vs fractal sleep cycle durations on the one side and either the amplitude of fractal descend/ascend (to reflect fractal cycle depth), duration of cycles with skipped REM sleep, duration of wake after sleep onset or the REM episode length of a given cycle (to reflect peak flatness) on the other side (Table 3).

Likewise, we computed person-centered effect sizes, the approach that answers the question, “How many participants in the study showed the consistent with theoretical expectation effect?”. This approach helps to reveal data patterns that are missed by traditional statistical analyses (Grice et al., 2020). We calculated the sample prevalence by counting the number of significant correlations between fractal and classical cycle duration divided by the total number of cases (both significant and non-significant).

To assess the population prevalence of the findings with associated uncertainty, we used the Bayesian prevalence, accounting for the false positive rate of the statistical test (Ince et al., 2022). This method helps to estimate the proportion of the population that would show the effect if they were tested in this experiment or, in other words, the population within-participant replication probability (Ince et al., 2022). As an output, this method provides the maximum a posterior estimate – the most likely value of the population parameter. To quantify the uncertainty of this estimate, Bayesian prevalence also provides the highest posterior density intervals for various levels (we used the 96% probability level) – the range within which the true population value lies with the specified probability. To perform this analysis we used an online web application available at https://estimate.prevalence.online.

Earlier version: To compare pediatric and young adult groups (Table S3 in Supplementary Material 2), MDD patients and controls (Table 3), MDD patients treated with REM-suppressive antidepressants and patients treated with REM-non-suppressive antidepressants (Table S5 in Supplementary Material 2), we used the non-parametric Mann-Whitney U test. We performed the analyses both at the cycle level (while pooling the cycles of all participants together) as well as at the subject level (while averaging the cycles of a given participant). Given that the results of both analyses were similar, we report only the cycle level analysis for simplicity. To compare medicated and unmedicated states of the MDD patients (Table 3), we used the paired samples Wilcoxon test. Effect sizes were calculated with Cohen’s d.

Revised version: To compare pediatric and young adult groups (Table S3 in Supplementary Material), MDD patients and controls (Table 4), MDD patients treated with REM-suppressive antidepressants and patients treated with REM-non-suppressive antidepressants (Table S5 in Supplementary Material), we used the non-parametric Mann-Whitney U test. We performed the analyses both at the cycle level (while pooling the cycles of all participants together) as well as at the subject level (while averaging the cycles of a given participant). Given that the results of both analyses were similar, we report only the cycle level analysis for simplicity. To compare medicated and unmedicated states of the MDD patients (Table 4), we used the paired samples Wilcoxon test. Effect sizes were calculated with Cohen’s d.

Earlier version: In Supplementary Material 2, we report autocorrelations and partial autocorrelations of fractal slope time series (Fig.S11) as well as cross-correlations (Fig.S12) between time series of fractal slopes vs. time series of non-REM or REM sleep proportion to further model their temporal relationships.

Revised version: In Supplementary Material, we report autocorrelations and partial autocorrelations of fractal slope time series (Fig.S11) as well as cross-correlations (Fig.S12) between time series of fractal slopes vs. time series of non-REM or REM sleep proportion to further model their temporal relationships.

Data and code sharing

Earlier version: The fractal slope and sleep stage for each 30-second epoch of sleep for healthy adult participants, Matlab scripts calculating fractal slopes and fractal cycles and Excel file used to perform all the statistical analyses with the fractal and sleep characteristics for each participant can be accessed under https://osf.io/gxzyd.

Revised version: The fractal slope and sleep stage for each 30-second epoch of sleep for healthy adult participants, Matlab scripts calculating fractal slopes and fractal cycles, Excel file used to perform all the statistical analyses with the fractal and sleep characteristics for each participant and PowerPoint file depicting fractal and classical cycles for all participants can be accessed under https://osf.io/gxzyd.

Results

Fractal cycles in healthy adults

Earlier version: Fig.2 A displays smoothed fractal slope time series and hypnogram for an example subject. Four additional examples are presented in Fig.S2 (Supplementary Material 2). Fractal slope time series and hypnograms for all healthy adult participants are shown in Supplementary Material 1.

Revised version: Fig.2 A displays smoothed fractal slope time series and hypnogram for an example subject. Four additional examples are presented in Fig.S2 (Supplementary Material). Fractal slope time series and hypnograms for all healthy adult participants are shown in Supplementary PowerPoint File shared on https://osf.io/gxzyd.

We observed that the slopes of the fractal (aperiodic) power component fluctuate across a night such that the peaks of the time series largely coincide with REM sleep episodes while the troughs of the time series for the most part coincide with non-REM sleep episodes. Based on this observation we propose the following definition:

Definition: The fractal activity-based cycles of sleep or “fractal cycles” for short is a time interval during which the time series of the fractal slopes descend from the local maximum to the local minimum with the amplitudes higher than |0.9| z, and then lead back from that local minimum to the next local maximum.

Earlier version: Based on this definition, we created an algorithm, which automatically defined the onset and offset of the fractal cycles (the adjacent peaks of the time series of the fractal slopes) (https://osf.io/gxzyd/). An additional visual inspection showed that the automatic definition of fractal cycles (Fig.2 A, blue diamonds) was identical to that provided by a human scorer.

Revised version: Based on this definition, we created an algorithm, which automatically defined the onset and offset of the fractal cycles (the adjacent peaks of the time series of the fractal slopes) (available on https://osf.io/gxzyd/). An additional visual inspection showed that the automatic definition of fractal cycles (Fig.2 A, blue diamonds) was identical to that provided by a human scorer.

Earlier version: Overall, fractal slopes cyclically descend and ascend 4 – 6 times per night and the average duration of such a descent-ascent cycle is close to 90 minutes. Fig.S3 A (Supplementary Material 2) shows the frequency distribution of the fractal cycle durations for each dataset separately as well as for the pooled dataset.

Revised version: Overall, fractal slopes cyclically descend and ascend 4 – 6 times per night and the average duration of such a descent-ascent cycle is close to 90 minutes. Fig.S3 A (Supplementary Material) shows the frequency distribution of the fractal cycle durations for each dataset separately as well as for the pooled dataset.

Earlier version: This observation strikingly resembles what we know about classical sleep cycles: “night sleep consists of 4 – 6 sleep cycles, which last for about 90 minutes each” (Feinberg & Floid, 1979; Le Bon, 2020; Fig.S3 A, bottom panel). Further calculations showed that the mean duration of the fractal cycles averaged over all cycles from all datasets (n = 940) is 89 ± 34 minutes while the mean duration of the classical sleep cycles is 90 ± 25 minutes (Fig.S3 B, Supplementary Material 2). The mean durations of the fractal and classical sleep cycles averaged over each participant correlated in all analyzed datasets (r = 0.4 – 0.5, Table 2, Fig.2 C).

Revised version: This observation strikingly resembles what we know about classical sleep cycles: “night sleep consists of 4 – 6 sleep cycles, which last for about 90 minutes each” (Feinberg & Floid, 1979; Le Bon, 2020; Fig.S3 A, bottom panel). Further calculations showed that the mean duration of the fractal cycles averaged over all cycles from all datasets (n = 940) is 89 ± 34 minutes while the mean duration of the classical sleep cycles is 90 ± 25 minutes (Fig.S3 B, Supplementary Material). The mean durations of the fractal and classical sleep cycles averaged over each participant correlated in all analyzed datasets (r = 0.4 – 0.5, Table 2, Fig.2 C).

Earlier version: Fig.S7 (Supplementary Material 2) shows fractal activity across 13-h, including 3 hours before the sleep onset and 2 hours after awakening. The pattern of fractal fluctuations suggests that fractal cycles are specific to sleep and are not observed during wake.

Revised version: Fig.S7 (Supplementary Material) shows fractal activity across 13-h, including 3 hours before the sleep onset and 2 hours after awakening. The pattern of fractal fluctuations suggests that fractal cycles are specific to sleep and are not observed during wake.

Cycle-to-cycle overnight dynamics showed an inverted U-shape of the fractal cycle durations and a gradual decrease in absolute amplitudes of the fractal descents and ascents from early to late cycles. This pattern resembled an inverted U-shape of the classical cycle durations (Fig.2 D).

Correspondence between fractal and classical cycles

Analysis at the individual cycle level revealed that 81% (763/940) of all fractal cycles (77 – 88% in different datasets) could be matched to a specific classical cycle defined by hypnogram, i.e., the timings of fractal and classical cycles approximately coincide. Bayesian prevalence analysis further revealed that the Bayesian highest posterior density interval with 96% probability level lies within the 0.77 – 0.83 range (the range within which the true population value lies) and the maximum a posteriori point estimate prevalence is equal to 0.8, reflecting the most likely values for the population parameter. This analysis reflects the within-participant replication probability: the probability of obtaining a significant experimental result if the same experiment was applied to a new participant randomly selected from the population (Ince et al., 2022).

Earlier version: In 54% (111/205) of the participants (45 – 66% in different datasets), all fractal cycles approximately coincided with classical cycles (r = 0.5 – 0.8, p < 0.001, Table 2 and Fig.S4, Supplementary Material 2). Bayesian prevalence analysis revealed that the maximum a posteriori point estimate prevalence is equal to 0.52 and the Bayesian highest posterior density interval (the true population level) with 96% probability level lies within the 0.45 – 0.60 range.

Revised version: In 54% (111/205) of the participants (45 – 66% in different datasets), all fractal cycles approximately coincided with classical cycles (r = 0.5 – 0.8, p < 0.001, Table 2 and Fig.S4, Supplementary Material). Bayesian prevalence analysis revealed that the maximum a posteriori point estimate prevalence is equal to 0.52 and the Bayesian highest posterior density interval (the true population level) with 96% probability level lies within the 0.45 – 0.60 range.

Earlier version: In the remaining 46% of the participants, the difference between the fractal and classical cycle numbers ranged from -2 to 2 with the average of -0.23 ± 1.23 cycle. This subgroup had 4.6 ± 1.2 fractal cycles per participant, while the number of classical cycles was 4.9 ± 0.7 cycles per participant. The correlation coefficient between the fractal and classical cycle numbers was 0.280 (p = 0.006) and between the cycle durations – 0.278 (p = 0.006). Still, in these participants, many – even though not all – fractal cycles could be matched to a specific classical cycle. Fig.2 B displays such an example in one participant. More examples can be found in Fig.S2 C – D of Supplementary Material 2 and Supplementary Material 1.

Revised version: In the remaining 46% of the participants, the difference between the fractal and classical cycle numbers ranged from −2 to 2 with the average of −0.23 ± 1.23 cycle. This subgroup had 4.6 ± 1.2 fractal cycles per participant, while the number of classical cycles was 4.9 ± 0.7 cycles per participant. The correlation coefficient between the fractal and classical cycle numbers was 0.280 (p = 0.006) and between the cycle durations – 0.278 (p = 0.006). Still, in these participants, many – even though not all – fractal cycles could be matched to a specific classical cycle. Fig.2 B displays such an example in one participant. More examples can be found in Fig.S2 C – D of Supplementary Material and Supplementary PowerPoint File shared on https://osf.io/gxzyd.

Sources of fractal and classical cycle mismatches

Earlier version: The timings and correlations between the fractal and classical cycles were not one-to-one (r = 0.6 – 0.8, p < 0.001). We identified two possible sources of a mismatch (see also Table 5).

Revised version: The timings and correlations between the fractal and classical cycles were not one-to-one (r = 0.6 – 0.8, p < 0.001). We identified two possible sources of a mismatch (Table 3; see also Table 6).

Earlier version: REM episode duration. While the fractal cycle end is defined as the local maximum of time series of fractal slopes, the classical cycle ends with the end of a REM episode. As a consequence, in some cases, especially for morning cycles that have rather long REM periods (> 20 minutes), the match between fractal and classical cycles can be rather coarse-grained (See, for example, cycle 3 in S16, Fig.S2 A, Supplementary Material 2). Yet, in other cases, the match between fractal and classical cycles might be almost perfect (See Fig.2 A).

Revised version: 1) REM episode duration. While the fractal cycle end is defined as the local maximum of time series of fractal slopes, the classical cycle ends with the end of a REM episode. As a consequence, in some cases, especially for morning cycles that have rather long REM periods (> 20 minutes), the match between fractal and classical cycles can be rather coarse-grained (See, for example, cycle 3 in S16, Fig.S2 A, Supplementary Material). Yet, in other cases, the match between fractal and classical cycles might be almost perfect (See Fig.2 A).

Earlier version: To test this visual observation, we correlated the absolute values of the difference in classical vs fractal sleep cycle durations with the REM episode length within a given cycle. We included in this analysis only the participants who had an equal number of fractal and classical cycles in order to match each fractal cycle to a classical cycle individually. We found that longer REM episodes were associated with a higher difference between classical vs fractal sleep cycle durations (r = 0.36, p < 0.001, n = 417 cycles, Table S10 in Supplementary Material 2).

Revised version: To test this visual observation, we correlated the absolute values of the difference in classical vs fractal sleep cycle durations with the REM episode length within a given cycle. We included in this analysis only the participants who had an equal number of fractal and classical cycles in order to match each fractal cycle to a classical cycle individually. We found that longer REM episodes were associated with a higher difference between classical vs fractal sleep cycle durations (r = 0.36, p < 0.001, n = 417 cycles, Table 3).

Earlier version: Wake after sleep onset (WASO) duration. Visual inspection of the data suggested that participants with more WASO often had more fractal than classical cycles. This might stem from the fact that both REM- and wake-related smoothed fractal slopes could be defined as local peaks (Fig.2 A – B, Fig.S1 B, Supplementary Material 2). More fractal peaks imply more fractal cycles and thus, possibly, more mismatches between the number and duration of classical and fractal cycles. To test this hypothesis, we correlated the average difference between the durations of classical and fractal cycles for each participant with the WASO proportion. We found that a higher difference in cycle durations was associated with a higher WASO proportion in 3/5 datasets (r’s = 0.36 – 0.49, p < 0.030) as well as in the merged dataset (r = 0.23, p = 0.001, n = 205 participants, Supplementary Material 2, Table S10).

Revised version: 2) Wake after sleep onset (WASO) duration. Visual inspection of the data suggested that participants with more WASO often had more fractal than classical cycles. This might stem from the fact that both REM- and wake-related smoothed fractal slopes could be defined as local peaks (Fig.2 A – B, Fig.S1 B, Supplementary Material). More fractal peaks imply more fractal cycles and thus, possibly, more mismatches between the number and duration of classical and fractal cycles. To test this hypothesis, we correlated the average difference between the durations of classical and fractal cycles for each participant with the WASO proportion. We found that a higher difference in cycle durations was associated with a higher WASO proportion in 3/5 datasets (r’s = 0.36 – 0.49, p < 0.030) as well as in the merged dataset (r = 0.23, p = 0.001, n = 205 participants, Table 3).

Earlier version: In addition, we correlated the difference in classical vs fractal cycle durations with the fractal descent or ascent amplitudes (as reflections of fractal cycle depth and possibly sleep quality). We found that a shallower fractal descent was associated with a higher mismatch between fractal and classical cycles in 1/5 datasets (r = 0.33, p = 0.02) as well as in the merged dataset (r = 0.15, p = 0.002, n = 400 cycles, Supplementary Material 2, Table S10).

Revised version: In addition, we correlated the difference in classical vs fractal cycle durations with the fractal descent or ascent amplitudes (as reflections of fractal cycle depth and possibly sleep quality). We found that a shallower fractal descent was associated with a higher mismatch between fractal and classical cycles in 1/5 datasets (r = 0.33, p = 0.02) as well as in the merged dataset (r = 0.15, p = 0.002, n = 400 cycles, Table 3).

Fractal cycles in children and adolescents

Earlier version: Next, we explored fractal cycles in children and adolescents (mean age: 12.4 ± 3.1 years, n = 21, Table S3 of Supplementary Material 2) and compared them with those in young adults (mean age: 24.8 ± 0.9 years, n = 24). Two examples of smoothed fractal slope time series and hypnograms from the pediatric dataset are shown in Fig.3 A – B. All examples are shown in Supplementary Material 1.

Revised version: Next, we explored fractal cycles in children and adolescents (mean age: 12.4 ± 3.1 years, n = 21, Table S3 of Supplementary Material) and compared them with those in young adults (mean age: 24.8 ± 0.9 years, n = 24). Two examples of smoothed fractal slope time series and hypnograms from the pediatric dataset are shown in Fig.3 A – B. All examples are shown in Supplementary PowerPoint File shared on https://osf.io/gxzyd.

Fractal cycles in children and adolescents.

Earlier version: A – B: Individual cycles: time series of smoothed z-normalized fractal slopes (bottom) and corresponding hypnograms (top). The duration of the fractal cycle is a time interval between two successive peaks (blue diamonds) defined with the Matlab function findpeaks with a minimum peak distance of 20 minutes and minimum peak prominence of 0.9 z. SWS – slow-wave sleep, REM – rapid eye movement sleep. A: In this 9.9-year-old participant (from Dataset 6), we split the first 150-minute-long classical cycle into two cycles according to the definitions of a “skipped” cycle presented in Methods. The fractal cycle algorithm successfully detected this skipped cycle. B: This 14.9-year-old participant has a 156-minute-long first classical cycle. Visual inspection shows that it should be divided into 3 skipped cycles, however, our a priori definition of skipped cycles did not include an option to subdivide a long cycle into three short cycles; hence, we split it into two short cycles. The fractal cycle algorithm was sensitive to these sleep lightenings and detected all three short cycles. Classical cycle 4 looks like a skipped cycle as it has two clear episodes of slow-wave sleep separated by non-REM stage 2. However, the length of this cycle is shorter than 110 min (the threshold defined a priori), therefore, we did not split the classical cycle 4 into two cycles. The fractal cycle algorithm was sensitive to this lightening of sleep and defined two fractal cycles during this period. C. Histograms: The frequency distribution of fractal (left) and classical (right) cycle durations in children and adolescents (mean age: 12.4 ± 3.1 years) compared to young adults (mean age: 24.8 ± 0.9 years). Kolmogorov-Smirnov’s test rejected the assumption that cycle duration comes from a standard normal distribution. D. Box plots: in each box, a vertical central line represents the median, the left and right edges of the box indicate the 25th and 75th percentiles, respectively, the whiskers extend to the most extreme data points not considered outliers, and a plus sign represents outliers. Children and adolescents show shorter fractal cycle duration compared to young adults. E. Overnight dynamics: cycle-to-cycle dynamics show that the first and the second fractal cycles are shorter in the pediatric compared to control group, * marks a statistically significant difference between the groups.

Revised version: A – B: Individual cycles: time series of smoothed z-normalized fractal slopes (bottom) and corresponding hypnograms (top). The duration of the fractal cycle is a time interval between two successive peaks (blue diamonds) defined with the Matlab findpeaks function with a minimum peak distance of 20 minutes and minumum peak prominence of 0.9 z. SWS – slow-wave sleep, REM – rapid eye movement sleep. A: In this 9.9-year-old participant (from Dataset 6), we split the first 150-minute-long classical cycle into two cycles according to the definitions of a “skipped” cycle presented in Methods. The fractal cycle algorithm successfully detected this skipped cycle. B: This 14.9-year-old participant has a 156-minute-long first classical cycle. Visual inspection shows that it should be divided into 3 skipped cycles, however, our a priori definition of skipped cycles did not include an option to subdivide a long cycle into three short cycles; hence, we split it into two short cycles. The fractal cycle algorithm was sensitive to these sleep lightenings and detected all three short cycles. Classical cycle 4 looks like a skipped cycle as it has two clear episodes of slow-wave sleep separated by non-REM stage 2. However, the length of this cycle is shorter than 110 min (the threshold defined), therefore, we did not split the classical cycle 4 into two cycles. The fractal cycle algorithm was sensitive to this lightening of sleep and defined two fractal cycles during this period. C. Histograms : The frequency distribution of fractal (left) and classical (right) cycle durations in children and adolescents (mean age: 12.4 ± 3.1 years) compared to young adults (mean age: 24.8 ± 0.9 years). Kolmogorov-Smirnov’s test rejected the assumption that cycle duration comes from a standard normal distribution. D. Box plots:in each box, a vertical central line represents the median, the left and right edges of the box indicate the 25th and 75th percentiles, respectively, the whiskers extend to the most extreme data points not considered outliers, and a plus sign represents outliers. Children and adolescents show shorter fractal cycle duration compared to young adults E. Overnight dynamics: cycle-to-cycle dynamics show that the first and the second fractal cycles are shorter in the pediatric compared to control group, * marks a statistically significant difference between the groups.

We found that children and adolescents had shorter fractal cycles compared to young adults with a medium effect size (76 ± 34 vs 94 ± 32 min, p < 0.001, Cohen’s d = −0.57, 112 vs 121 pooled cycles, 5.0 cycles/participant vs 4.4 cycles/participant, Fig.3 C – D, Table S3). Similarly, children and adolescents showed shorter classical cycles than young adults with a medium effect size (80 ± 23 vs 90 ± 22 min, p < 0.001, Cohen’s d = −0.42, 112 vs 114 pooled cycles, Fig.3 C – D).

To directly compare the fractal and classical approaches, we performed a Multivariate Analysis of Variance with fractal and classical cycle durations as dependent variables, the group as an independent variable and the age as a covariate. We found that fractal cycle durations showed higher F-values (F(1, 43) = 4.5 vs F(1, 43) = 3.1), adjusted R squared (0.138 vs 0.089) and effect sizes (partial eta squared 0.18 vs 0.13) than classical cycle durations.

Cycle-to-cycle overnight dynamics further revealed that the first and second fractal – but not classical – cycles were significantly shorter in the pediatric compared to the control group (Fig.3 E) with medium effect sizes (d = −0.61 – −0.72). At the same time, the overnight classical – but not fractal – cycle analysis detected a between-group difference for the fourth classical cycle with a large effect size (d = −1.0, Fig.3 E).

Skipped cycles

Earlier version: We tested whether the fractal cycle algorithm can detect skipped cycles, i.e., the cycles where an anticipated REM episode is skipped (possibly due to too high homeostatic non-REM pressure). We performed this analysis twice. First, we counted all skipped cycles (except the last cycles of a night, which usually lack a REM episode for other reasons, e.g., a participant had/was woken up). Second, we counted only the first classical cycles (i.e., the first cycle out of the 4 – 6 cycles that each participant had per night, Fig. 3 A – B) as these cycles coincide with the highest non-REM pressure. An additional reason to disregard skipped cycles observed later during the night was our aim to achieve higher between-subject consistency as later skipped cycles were observed in only a small number of participants.

Revised version: We tested whether the fractal cycle algorithm can detect skipped cycles, i.e., the cycles where an anticipated REM episode is skipped possibly due to too high homeostatic non-REM pressure. We counted only the first classical cycles (i.e., the first cycle out of the 4 – 6 cycles that each participant had per night, Fig. 3 A – B) as these cycles coincide with the highest non-REM pressure. An additional reason to disregard skipped cycles observed later during the night was our aim to achieve higher between-subject consistency as second – sixth skipped cycles were observed in only a small number of participants and were not distributed equally across the datasets.

Earlier version: The average number of the first skipped cycles for Datasets 1 – 5 is reported in Table 2. Table S9 of Supplementary Material 2 further reports the average number of first and all skipped cycles as assessed by two independent human raters and the inter-rater agreement. Three specific examples of skipped cycles in young adults are presented in Fig.S6 of Supplementary Material 2 and two examples in children are shown in Fig.3 A – B. All cycles can be found in Supplementary Material 1.

Revised version: The average number of the first skipped cycles for Datasets 1 – 5 is reported in Table 2. Table S9 of Supplementary Material further reports the average number of skipped cycles as assessed by two independent human raters and the inter-rater agreement. Three specific examples of skipped cycles in young adults are presented in Fig.S6 of Supplementary Material and two examples in children are shown in Fig.3 A – B. All cycles are marked in Supplementary PowerPoint File shared on https://osf.io/gxzyd.

Removed in the revised version: Visual inspection of the hypnograms from Datasets 1 – 6 revealed that 96 out of 843 (11%) sleep cycles lacked REM episodes. The fractal cycle algorithm detected skipped cycles in 90 out of 97 (93%) cases. Bayesian prevalence analysis revealed that the maximum a posteriori point estimate prevalence equaled 0.92. The Bayesian highest posterior density interval (the true population level) with a 96% probability level lay within the 0.85 – 0.97 range. For reference, we also assessed the agreement on the presence/absence of skipped cycles between two human raters and found that it equaled 61% (91 cycles detected by both raters out of 149 cycles detected by either one or two scorers); that is 32% lower than the performance of the fractal cycle algorithm.

Earlier version: Visual inspection of the hypnograms from Datasets 1 – 6 revealed that 57 out of the 226 (25%) first sleep cycles lacked REM episodes. The fractal cycle algorithm detected skipped cycles in 54 out of 57 (95%) cases. Bayesian maximum a posteriori point estimate prevalence equaled 0.96 with highest posterior density interval lying within the 0.88 – 0.99 range. The agreement between two human raters was 85% (55 cycles detected by both raters out of 65 cycles detected by either one or two scorers); that is 10% lower than the performance of the fractal cycle algorithm.

Revised version: Visual inspection of the hypnograms from Datasets 1 – 6 was performed by two independent researchers. Scorer 1 and Scorer 2 detected that out of 226 first sleep cycles 58 (26%) and 64 (28%), respectively, lacked REM episodes. The agreement on the presence of skipped cycles between two human raters equaled 91% (58 cycles detected by both raters out of 64 cycles detected by either one or two scorers). The fractal cycle algorithm detected skipped cycles in 57 out of 58 (98%) cases detected by Scorer 1 with one false positive (which, however, was tagged as a skipped cycle by Scorer2), and in 58 out of 64 (91%) cases detected by Scorer 2 with no false positives.

Age and fractal cycles

Earlier version: Next, we tested whether fractal cycle duration changes as a function of age. We found that in the merged adult dataset (Datasets 1 – 5, n = 205), the mean duration of the fractal cycles negatively correlated with the age of the participants (r = -0.19, p = 0.006, age range: 18 – 75 years, median: 33.5 years, Fig.S5 A, Supplementary Material 2). Intriguingly, this correlation looked like a mirror image of the correlation between the age and wakefulness after sleep onset (Fig.S5 B). Following this observation, we performed an additional correlation between the fractal cycle duration and wakefulness proportion and found that it was non-significant (r = 0.01, p = 0.969). Nevertheless, we further performed a partial correlation between the fractal cycle duration and participant age, while controlling for the effect of wakefulness after the sleep onset and found that the correlation remained significant (r = -0.18, p = 0.011).

Revised version: Next, we tested whether fractal cycle duration changes as a function of age. We found that in the merged adult dataset (Datasets 1 – 5, n = 205), the mean duration of the fractal cycles negatively correlated with the age of the participants (r = −0.19, p = 0.006, age range: 18 – 75 years, median: 33.5 years, Fig.S5 A, Supplementary Material). Intriguingly, this correlation looked like a mirror image of the correlation between the age and wakefulness after sleep onset (Fig.S5 B). Following this observation, we performed an additional correlation between the fractal cycle duration and wakefulness proportion and found that it was non-significant (r = 0.01, p = 0.969). Nevertheless, we further performed a partial correlation between the fractal cycle duration and participant age, while controlling for the effect of wakefulness after the sleep onset and found that the correlation remained significant (r = −0.18, p = 0.011).

Given that participant’s age also correlated with REM latency (Fig.S5 D) while REM latency further correlated with fractal cycle duration (Fig.S5 C), we performed an additional partial correlation between the fractal cycle duration and age while controlling for REM latency. We found that it remained significant (r = −0.16, p = 0.025). The partial correlation between the fractal cycle duration and REM latency adjusted for the participant’s age was non-significant (r = 0, p = 0.746).

Of note, these correlations were significant while analyzing the pooled dataset only, they were not observed while analyzing each dataset separately. Moreover, when we added to the pooled adult dataset (Datasets 1 – 5) our pediatric dataset (Dataset 6), the correlation between fractal cycle duration and age became non-significant.

Interestingly, the mean duration of the classical cycles did not correlate with the age of the adult participants neither in the merged dataset (r = −0.02, p = 0.751) nor while analyzing each dataset separately.

Fractal cycles in MDD

Earlier version: Finally, to assess the clinical relevance of the fractal cycles, we explored them in patients with MDD. We found that patients at 7- and 28-day of medication treatment as well as long-termed medicated patients (Datasets A – C) showed a longer fractal cycle duration compared to controls with medium effect size (Table 3, Fig.4 B). Moreover, in Dataset B, the patients who took REM-suppressive antidepressants (See Table S5 of Supplementary Material 2 for information on specific medications taken by the patients) showed longer fractal cycle duration compared to patients who took REM-non-suppressive antidepressants with medium effect size (70 cycles of 21 patients vs 63 cycles of 17 patients). In Dataset C, no difference was detected between these sub-groups. However, it should be noted that in Datasets C, the REM-suppressive and REM-non-suppressive antidepressant groups were unbalanced (87 cycles of 23 patients vs 35 cycles of 10 patients) and consisted of different medications than Dataset B.

Revised version: Finally, to assess the clinical relevance of the fractal cycles, we explored them in patients with MDD. We found that patients at 7- and 28-day of medication treatment as well as long-termed medicated patients (Datasets A – C) showed a longer fractal cycle duration compared to controls with medium effect size (Table 4, Fig.4 B). Moreover, in Dataset B, the patients who took REM-suppressive antidepressants (See Table S5 of Supplementary Material for information on specific medications taken by the patients) showed longer fractal cycle duration compared to patients who took REM-non-suppressive antidepressants with medium effect size (70 cycles of 21 patients vs 63 cycles of 17 patients). In Dataset C, no difference was detected between these sub-groups. However, it should be noted that in Datasets C, the REM-suppressive and REM-non-suppressive antidepressant groups were unbalanced (87 cycles of 23 patients vs 35 cycles of 10 patients) and consisted of different medications than Dataset B.

Fractal cycles in MDD.

Earlier version: A. Individual fractal cycles: time series of smoothed z-normalized fractal slopes observed in a 22 y.o. MDD patient (Dataset B) in their unmedicated (top) and 7-day medicated (bottom) states. Peaks (blue diamonds) are defined with the Matlab function findpeaks with the minimum peak distance of 20 minutes and minimum peak prominence of 0.9 z. Fractal cycles duration (defined as an interval of time between two successive peaks) is longer in the medicated compared to unmedicated states, reflecting shallower fluctuations of fractal (aperiodic) activity. Two additional patients are shown in Fig.S9 (Supplementary Material 2). B. Box plots: the fractal cycle duration is longer in medicated MDD patients (red) compared to age and gender-matched healthy controls (black) in all datasets. In Dataset B, fractal cycles are longer in the medicated vs patients’ own unmedicated state and in patients who took REM-suppressive vs REM-non-suppressive antidepressants. A vertical central line represents the median in each box, the left and right edges of the box indicate the 25th and 75th percentiles, respectively, the whiskers extend to the most extreme data points not considered outliers, and a plus sign represents outliers (individual cycles). C. Frequency distribution: individual fractal and classical cycles pooled from three MDD datasets (A – C) are counted separately for medicated MDD patients and HC. D. Overnight dynamics: cycle-to-cycle dynamics of the duration of both fractal and classical cycles show a gradual decrease in medicated patients vs an inverted U shape in controls. The between-group difference in cycle duration is the largest for the first cycle. Patients show flatter fractal descents of the second cycle and steeper fractal descents of the fourth cycle compared to controls. Contrary to controls, patients do not show a gradual decrease in absolute amplitudes of the fractal descents from the second to the fourth cycles. Patients and controls show comparable cycle-to-cycle dynamics of fractal ascents, * marks a statistically significant difference between the groups. MDD – major depressive disorder, HC – healthy controls, unmed. – unmedicated, med. – medicated, SWS – slow-wave sleep, REM – rapid eye movement.

Revised version: A. Individual fractal cycles: time series of smoothed z-normalized fractal slopes observed in a 22 y.o. MDD patient (Dataset B) in their unmedicated (top) and 7-day medicated (bottom) states. Peaks (blue diamonds) are defined with the Matlab function findpeaks with the minimum peak distance of 20 minutes and minimum peak prominence of 0.9 z. Fractal cycles duration (defined as an interval of time between two successive peaks) is longer in the medicated compared to unmedicated states, reflecting shallower fluctuations of fractal (aperiodic) activity. Two additional patients are shown in Fig.S9 (Supplementary Material). B. Box plots: the fractal cycle duration is longer in medicated MDD patients (red) compared to age and gender-matched healthy controls (black) in all datasets. In Dataset B, fractal cycles are longer in the medicated vs patients’ own unmedicated state and in patients who took REM-suppressive vs REM-non-suppressive antidepressants. A vertical central line represents the median in each box, the left and right edges of the box indicate the 25th and 75th percentiles, respectively, the whiskers extend to the most extreme data points not considered outliers, and a plus sign represents outliers (individual cycles). C. Frequency distribution: individual fractal and classical cycles pooled from three MDD datasets (A – C) are counted separately for medicated MDD patients and HC. D. Overnight dynamics: cycle-to-cycle dynamics of the duration of both fractal and classical cycles show a gradual decrease in medicated patients vs an inverted U shape in controls. The between-group difference in cycle duration is the largest for the first cycle. Patients show flatter fractal descents of the second cycle and steeper fractal descents of the fourth cycle compared to controls. Contrary to controls, patients do not show a gradual decrease in absolute amplitudes of the fractal descents from the second to the fourth cycles. Patients and controls show comparable cycle-to-cycle dynamics of fractal ascents, * marks a statistically significant difference between the groups. MDD – major depressive disorder, HC – healthy controls, unmed. – unmedicated, med. – medicated, SWS – slow-wave sleep, REM – rapid eye movement.

Earlier version: Table 3 and Fig.4 show results calculated over frontal electrodes (or central ones for Dataset A). The topographical analysis over other areas is reported in Table S6 of Supplementary Material 2.

Revised version: Table 4 and Fig.4 show results calculated over frontal electrodes (or central ones for Dataset A). The topographical analysis over other areas is reported in Table S6 of Supplementary Material.

Earlier version: In Dataset B (the only dataset including unmedicated patients), 7-day medicated patients had longer fractal cycles compared to their own unmedicated state with medium effect size (p = 0.001, Cohen’s d = 0.4, Fig.4 A – B, two additional examples are shown in Fig.S10, Supplementary Material 2). Unmedicated patients and controls showed comparable durations of the fractal cycles. The only difference observed between these groups was a smaller amplitude of the fractal descent of the first fractal cycles in unmedicated patients compared to controls with a medium effect size (−3.2 vs -3.6 z, p = 0.040, Cohen’s d = 0.5).

Revised version: In Dataset B (the only dataset including unmedicated patients), 7-day medicated patients had longer fractal cycles compared to their own unmedicated state with medium effect size (p = 0.001, Cohen’s d = 0.4, Fig.4 A – B, two additional examples are shown in Fig.S10, Supplementary Material). Unmedicated patients and controls showed comparable durations of the fractal cycles. The only difference observed between these groups was a smaller amplitude of the fractal descent of the first fractal cycles in unmedicated patients compared to controls with a medium effect size (−3.2 vs −3.6 z, p = 0.040, Cohen’s d = 0.5).

In a pooled dataset, medicated patients showed a prolonged duration of fractal cycles compared to the controls (104 ± 49 vs 88 ± 31 min, p < 0.001, Fig.4 C). The between-group difference was the largest for the first cycle (Fig.4 D). Moreover, cycle-to-cycle overnight dynamics of the fractal cycle duration showed a gradual decrease in medicated patients vs an inverted U shape in controls (Fig.4 D).

Earlier version: To test our hypothesis that fractal cycles are more sensitive than classical cycles in detecting differences between patients and controls, we performed the same analysis as described above while using the duration of classical cycles as the variable of interest. The results were similar to those obtained for fractal cycle durations (Table 3, Fig.4 C – D), i.e., our hypothesis was not confirmed. The comparable outcomes of the two analyses can be explained by the positive correlations between the durations of fractal and classical cycles observed in all groups of the medicated MDD patients like that seen in healthy controls (Table 3).

Revised version: To test our hypothesis that fractal cycles are more sensitive than classical cycles in detecting differences between patients and controls, we performed the same analysis as described above while using the duration of classical cycles as the variable of interest. The results were similar to those obtained for fractal cycle durations (Table 4, Fig.4 C – D), i.e., our hypothesis was not confirmed. The comparable outcomes of the two analyses can be explained by the positive correlations between the durations of fractal and classical cycles observed in all groups of the medicated MDD patients like that seen in healthy controls (Table 4).

Discussion

This study introduced the new concept of fractal activity-based cycles of sleep or “fractal cycles” for short, which is based on temporal fluctuations of the fractal (aperiodic) slopes across a night. We showed that durations of these fractal cycles correlated with those of classical (non-REM – REM) sleep cycles defined by hypnograms in five independently collected datasets counting 205 healthy participants overall as well as in 111 medicated patients with MDD. Overnight cycle-to-cycle dynamics in healthy adults showed an inverted U-shape for both fractal and classical cycle durations. The fractal cycle algorithm was effective in detecting cycles with skipped REM sleep. The findings further revealed that children and adolescents showed shorter fractal cycles as compared to young healthy adults. In adults, fractal cycle durations negatively correlated with participants’ age. Medicated patients with MDD showed longer fractal cycles compared to their own unmedicated state and healthy controls. Below we discuss these findings in detail.

Fractal cycles: definition and motivation

We observed that the time series of fractal slopes have a cyclical nature, descending and ascending for about 4 – 6 times per night with a mean duration of approximately 90 minutes for each such (“fractal”) cycle. This strikingly resembles the description of classical sleep cycles. Indeed, both the visual inspection and formal correlational analyses revealed that the timing and duration of the fractal and classical cycles mainly matched. This led us to propose that the “fractal cycles of sleep” could serve as a new data-driven definition of sleep cycles, i.e., a means to appreciate quantitatively what has been previously observed only qualitatively using hypnograms. Notably, we do not claim that fractal cycles are a substitute for the study of the individual sleep stages or microstructural features of sleep. We want to stress, however, that currently, sleep research is shifted towards the study of, to use a metaphor, “the atoms” of sleep, such as individual sleep stages, slow oscillations, spindles, microarousals etc. Yet it is possible that some important (currently unknown) features of sleep could be explored only at the level of sleep cycles, “the molecules of sleep”. (Note, that we use the molecule and atom concepts only as a metaphor for the macro- and microstructure of sleep.)

Hypothetical functional significance of aperiodic activity and fractal cycles

Earlier version: The decision to incorporate fractal activity analysis in sleep cycle research was based on the reports that fractal (aperiodic) dynamics may reflect the bistability of the network (the overall tendency of alternating up and down states) (Baranauskas et al., 2012) and/or alterations in the balance between neural excitatory and inhibitory currents (Gao et al., 2017). Circumstantial evidence suggests that fractal activity is a measure of sleep homeostasis or sleep intensity, reflecting sleep-wake history, sleep stage differences, sleep cycles, age-effects, local sleep and sleep disorders (Bódizs et al., 2023). Recently, it has been reported that during human sleep, spectral slopes positively correlate with pupil size, a marker of arousal levels linked to the activity of the locus coeruleus-noradrenergic system (Carro-Domínguez et al., 2023).

Revised version: The decision to incorporate fractal activity analysis in sleep cycle research was based on the reports that fractal (aperiodic) dynamics may reflect the bistability of the network (the overall tendency of alternating up and down states) (Baranauskas et al., 2012) and/or alterations in the balance between neural excitatory and inhibitory currents (Gao et al., 2017). Circumstantial evidence suggests that fractal activity is a measure of sleep homeostasis or sleep intensity, reflecting sleep-wake history, sleep stage differences, sleep cycles, age-effects, local sleep and sleep disorders (Bódizs et al., 2024). Recently, it has been reported that during human sleep, spectral slopes positively correlate with pupil size, a marker of arousal levels linked to the activity of the locus coeruleus-noradrenergic system (Carro-Domínguez et al., 2023).

According to the reciprocal-interaction model of sleep cycles, each sleep phase is characterized by a specific neurochemical mixture. During non-REM sleep, aminergic inhibition decreases and cholinergic excitation increases such that at REM sleep onset, aminergic inhibition is shut off and cholinergic excitability reaches its maximum, while other outputs are inhibited (Pace-Schott & Hobson, 2002). Complex inhibitory and excitatory connections between pontine REM-on and REM-off neurons are further modulated by such neurotransmitters as GABA, glutamate, nitric oxide and histamine. Intriguingly, during REM sleep, acetylcholine plays the main role in maintaining brain activation, which is expressed as EEG desynchronization, one of the main features of REM sleep, and other systems are silent (Nir & Tononi, 2010). This suggests that acetylcholine, which fluctuates cyclically across a night as a result of the REM-off – REM-on interactions, might have a key role in the sleep phase alternation.

Earlier version: Given that the specific neurochemical milieu of the brain produces a specific type of conscious experience (Nir & Tononi, 2010) and that conscious experience was shown to be related to fractal activity derived from the human sleep EEG (Colombo et al., 2019), it is tempting to speculate that fractal activity tracks sleep-related changes in the neurochemical milieu of the brain and overall network dynamics. This has not been tested in humans; nevertheless, in rats, cholinergic nucleus basalis stimulation acutely increased higher to lower frequency cortical LFP power ratio or in other words, caused flattering of spectral decay (Goard & Dan, 2009). One can, therefore, speculate that ascending parts and peaks of fractal cycles coincide with acetylcholine release. The troughs of fractal cycles, in turn, might reflect a higher homeostatic pressure and even cause feelings of sleepiness and the search for the opportunity of initiating sleep, as these are periods of the steepest fractal activity, which implies a higher ratio of lower over higher frequency power in the EEG (Bódizs et al., 2023).

Revised version: Given that the specific neurochemical milieu of the brain produces a specific type of conscious experience (Nir & Tononi, 2010) and that conscious experience was shown to be related to fractal activity derived from the human sleep EEG (Colombo et al., 2019), it is tempting to speculate that fractal activity tracks sleep-related changes in the neurochemical milieu of the brain and overall network dynamics. This has not been tested in humans; nevertheless, in rats, cholinergic nucleus basalis stimulation acutely increased higher to lower frequency cortical LFP power ratio or in other words, caused flattering of spectral decay (Goard & Dan, 2009). One can, therefore, speculate that ascending parts and peaks of fractal cycles coincide with acetylcholine release. The troughs of fractal cycles, in turn, might reflect a higher homeostatic pressure and even cause feelings of sleepiness and the search for the opportunity of initiating sleep, as these are periods of the steepest fractal activity, which implies a higher ratio of lower over higher frequency power in the EEG (Bódizs et al., 2024).

Earlier version: In view of this literature, we speculate that fractal fluctuations may reflect two antagonistic roles of sleep (Simor et al., 2022). Specifically, fractal cycle troughs might cohere with sensory disconnection that facilitates restorative properties of sleep while fractal cycle peaks reflect monitoring of the environment that transiently restores alertness (Table 4).

Revised version: In view of this literature, we speculate that fractal fluctuations may reflect two antagonistic roles of sleep (Simor et al., 2022). Specifically, fractal cycle troughs might cohere with sensory disconnection that facilitates restorative properties of sleep while fractal cycle peaks reflect monitoring of the environment that transiently restores alertness (Table 5).

Earlier version: Fractal and classical cycles comparison (Table 5)

Revised version: Fractal and classical cycles comparison (Table 6)

In this study, in healthy adults, 81% of all fractal cycles defined by our algorithm could be matched to individual classical cycles defined by hypnograms. Correlations between the durations of fractal and classical cycles were observed not only in healthy adults but also in MDD patients who took antidepressants. The results show that displaying sleep data using fractal activity as a function of time meaningfully adds to the conventionally used hypnograms thanks to the gradual and objective quality of fractal power.

Earlier version: Thus, in hypnograms, each sleep stage is ascribed with a rather arbitrary assigned categorical value (e.g., wake = 0, REM = -1, N1 = -2, N2 = -3 and SWS = -4, Fig.2 A). Yet categorical labeling of sleep stages can induce information loss and lead to several misinterpretations, such as an implied order of sleep stages (e.g., “REM sleep is located between wake and N1”) and an implied “attractor state” conception of sleep stages (e.g., “no inter-stage states”). Likewise, defining the precise beginning and end of a classical sleep cycle using a hypnogram is often difficult and arbitrary, for example, in cycles with skipped or interrupted REM sleep or REM sleep without atonia.

Revised version: Thus, in hypnograms, each sleep stage is ascribed with a categorical value (e.g., wake = 0, REM = −1, N1 = −2, N2 = −3 and SWS = −4, Fig.2 A). Yet categorical labeling of sleep stages can induce information loss and lead to several misinterpretations, such as an implied order of sleep stages (e.g., “REM sleep is located between wake and N1”) and an implied “attractor state” conception of sleep stages (e.g., “no inter-stage states”). Likewise, defining the precise beginning and end of a classical sleep cycle using a hypnogram is often difficult and arbitrary, for example, in cycles with skipped or interrupted REM sleep or REM sleep without atonia.

In contrast, fractal cycles do not rely on the assignment of categories, being based on a real-valued metric with known neurophysiological functional significance. This introduces a biological foundation and a more gradual impression of nocturnal changes compared to the abrupt changes that are inherent to hypnograms.

Earlier version: Importantly, fractal cycle computation is automatic and thus objective. Even though recently, there has been a significant surge in sleep analysis incorporating various machine learning techniques and deep neural network architectures, we should stress that this research line mainly focused on the automatic classification of sleep stages and disorders almost ignoring the area of sleep cycles. Here, as the reference method, we used one of the very few available algorithms for sleep cycle detection (Blume & Cajochen, 2021). We found that automatically identified classical sleep cycles only moderately correlated with those detected by human raters (r’s = 0.3 – 0.7 in different datasets). These coefficients lay within the range of the coefficients between fractal and classical cycle durations (r = 0.41 – 0.55, moderate) and outside the range of the coefficients between classical cycle durations detected by two human scorers (r’s = 0.7 – 0.9, strong, Supplementary Material 2, Table S8).

Revised version: Importantly, fractal cycle computation is automatic and thus objective. Even though recently, there has been a significant surge in sleep analysis incorporating various machine learning techniques and deep neural network architectures, we should stress that this research line mainly focused on the automatic classification of sleep stages and disorders almost ignoring the area of sleep cycles. Here, as the reference method, we used one of the very few available algorithms for sleep cycle detection (Blume & Cajochen, 2021). We found that automatically identified classical sleep cycles only moderately correlated with those detected by human raters (r’s = 0.3 – 0.7 in different datasets). These coefficients lay within the range of the coefficients between fractal and classical cycle durations (r = 0.41 – 0.55, moderate) and outside the range of the coefficients between classical cycle durations detected by two human scorers (r’s = 0.7 – 0.9, strong, Supplementary Material, Table S8).

Earlier version: One of the most significant methodological strengths of the fractal cycle algorithm is its ability to detect cycles with skipped REM sleep common in children, adolescents and young adults. Our algorithm detected skipped cycles in 93% of cases while the hypnogram-based agreement on the presence/absence of skipped cycles between two independent human raters was 61% only; thus, 32% lower. We deduce that the fractal cycle algorithm detected skipped cycles since a lightening of sleep that replaces a REM episode in skipped cycles is often expressed as a local peak in fractal time series. Based on this, we further hypothesize that, analogously, fractal cycles might detect REM sleep without atonia episodes in REM sleep behaviour disorder, the episodes currently often mistaken as non-REM sleep.

Revised version: One of the most significant methodological strengths of the fractal cycle algorithm is its ability to detect cycles with skipped REM sleep common in children, adolescents and young adults. Our algorithm detected skipped cycles in 91 – 98% of cases. We deduce that the fractal cycle algorithm detected skipped cycles since a lightening of sleep that replaces a REM episode in skipped cycles is often expressed as a local peak in fractal slope time series. Based on this, we further hypothesize that, analogously, fractal cycles might detect REM sleep without atonia episodes in REM sleep behaviour disorder, the episodes currently often mistaken as non-REM sleep.

In summary, we expect that fractal cycles could bring insights into (yet) unexplained phenomena thanks to their gradual and objective quality, and, therefore, have the potential to induce a paradigm shift in basic and clinical (see below) sleep research.

Fractal slopes and SWA: overnight dynamics

Earlier version: Of note, currently, the gold standard marker of many sleep functions (e.g., restorative, regenerative) with a long-standing use is slow-wave activity (SWA), which, similar to fractal slopes, is also continuous and objective. SWA, however, has several disadvantages, such as large variability between individuals, which makes it impossible to set up a given reference point for healthy sleep (Horváth et al., 2022). Interindividual variability of spectral slopes is much smaller compared to SWA, making it a less individual-specific metric, yet spectral slopes strongly correlate with SWA (31 – 53% of shared variance throughout the non-REM periods) (Horváth et al., 2022; Bódizs et al., 2023). In addition, both the literature and our findings show that while SWA has a cycling nature during the first part of the night, neural dynamics of late-night’s sleep are not reflected by SWA at all (Fig.S8 in Supplementary Material 2). Given that SWA is a primary marker of sleep homeostasis, this pattern possibly reflects the dissipation of a sleep need over the night (Bódizs et al., 2023). In contrast, fractal slopes show a cycling nature over the entire night’s sleep (Fig.2 A – B and Fig.S8), suggesting that they are a more suitable means to reflect the macrostructure of the whole night’s sleep than SWA.

Revised version: Of note, currently, the gold standard marker of many sleep functions (e.g., restorative, regenerative) with a long-standing use is slow-wave activity (SWA), which, similar to fractal slopes, is also continuous and objective. SWA, however, has several disadvantages, such as large variability between individuals, which makes it impossible to set up a given reference point for healthy sleep (Horváth et al., 2022). Interindividual variability of spectral slopes is much smaller compared to SWA, making it a less individual-specific metric, yet spectral slopes strongly correlate with SWA (31 – 53% of shared variance throughout the non-REM periods) (Horváth et al., 2022; Bódizs et al., 2024). In addition, both the literature and our findings show that while SWA has a cycling nature during the first part of the night, neural dynamics of late-night’s sleep are not reflected by SWA at all (Fig.S8 in Supplementary Material). Given that SWA is a primary marker of sleep homeostasis, this pattern possibly reflects the dissipation of a sleep need over the night (Bódizs et al., 2024). In contrast, fractal slopes show a cycling nature over the entire night’s sleep (Fig.2 A – B and Fig.S8), suggesting that they are a more suitable means to reflect the macrostructure of the whole night’s sleep than SWA.

Earlier version: Having said this, we should highlight that characteristics of fractal cycles of sleep do undergo some overnight changes. Thus, the durations of both fractal and classical cycles in health show an inverted U-shape across a night and the amplitudes of fractal descents and ascents are larger during early-night-compared to late-night cycles (Fig.2 D). This is in line with the report on the flattening of fractal activity from early to late sleep cycles (Horváth et al., 2023). If seen in the context of the reactive and predictive homeostatic functions of sleep (Simor et al., 2023), deeper fractal cycles observed during early-night sleep could reflect intensive restorative processes (which are also reflected by SWA), whereas shallower fractal cycles seen during the later part of night’s sleep could reflect more active future-oriented processes (which are not reflected by SWA) with a shift towards neural excitation relative to inhibition expressed as overall flatter fractal activity (Table 4).

Revised version: Having said this, we should highlight that characteristics of fractal cycles of sleep do undergo some overnight changes. Thus, the durations of both fractal and classical cycles in health show an inverted U-shape across a night and the amplitudes of fractal descents and ascents are larger during early-night-compared to late-night cycles (Fig.2 D). This is in line with the report on the flattening of fractal activity from early to late sleep cycles (Horváth et al., 2023). If seen in the context of the reactive and predictive homeostatic functions of sleep (Simor et al., 2023), deeper fractal cycles observed during early-night sleep could reflect intensive restorative processes (which are also reflected by SWA), whereas shallower fractal cycles seen during the later part of night’s sleep could reflect more active future-oriented processes (which are not reflected by SWA) with a shift towards neural excitation relative to inhibition expressed as overall flatter fractal activity (Table 5).

Fractal cycles and age

Earlier version: We found that older healthy participants had shorter fractal cycles compared to the younger ones while classical cycles did not correlate with the participants’ age. At first glance, it looked as if this association simply reflected an increased proportion of the wake after the sleep onset often seen in older adults (Fig.S5 B, Supplementary Material 2). Indeed, our algorithm does not discriminate between the smoothened wake- and REM-related fractal slopes and can define both as local peaks (Fig.2 A – B). This happens because for the most part, wake- and REM sleep-related smoothed fractal slopes display comparable values, which are also the highest ones compared to other stages (Fig.S1 B, green squares, Supplementary Material 2). Since the fractal cycle duration is defined as an interval of time between two adjacent peaks, more awakenings during sleep are expected to result in more peaks and, consequently, more fractal cycles per total sleep time, i.e., a shorter cycle duration. (It is worth mentioning that unsmoothed wake- and REM-related slopes differ (Schneider et al., 2022 and Fig.S1 B here (black squares). However, this is a side notion as raw values were not used in this study since our algorithm performed poorly on raw time series).

Revised version: We found that older healthy participants had shorter fractal cycles compared to the younger ones while classical cycles did not correlate with the participants’ age. At first glance, it looked as if this association simply reflected an increased proportion of the wake after the sleep onset often seen in older adults (Fig.S5 B, Supplementary Material). Indeed, our algorithm does not discriminate between the smoothened wake- and REM-related fractal slopes and can define both as local peaks (Fig.2 A – B). This happens because for the most part, wake- and REM sleep-related smoothed fractal slopes display comparable values, which are also the highest ones compared to other stages (Fig.S1 B, green squares, Supplementary Material). Since the fractal cycle duration is defined as an interval of time between two adjacent peaks, more awakenings during sleep are expected to result in more peaks and, consequently, more fractal cycles per total sleep time, i.e., a shorter cycle duration. (It is worth mentioning that unsmoothed wake- and REM-related slopes differ (Schneider et al., 2022 and Fig.S1 B here (black squares). However, this is a side notion as raw values were not used in this study since our algorithm performed poorly on raw time series).

Earlier version: Moreover, a larger difference in classical vs fractal cycle duration was associated with a higher proportion of wake after sleep onset (WASO) in 3/5 datasets as well as in the merged dataset (Supplementary Material 2, Table S10). On the other hand, the partial correlation between fractal cycle duration and age remained significant after controlling for the WASO amount. This hints that the association between fractal cycles and age might reflect more than just a confounding effect of WASO. This interpretation is in line with literature on age-related changes in aperiodic activity, namely, on flattering of fractal slopes with age (Voytek et al., 2015; Bódizs et al., 2021; Pathania et al., 2022), especially during SWS (Schneider et al., 2022). Likewise, aging is associated with shorter and fewer classical cycles, with a mean of 3.5 cycles per night compared to the usual 4 – 5 in adults and adolescents (Conte et al., 2014). Our findings suggest that fractal cycles are more sensitive to these age-related alterations than the classical ones. We further speculate that the claim that “age affects sleep microstructure more than sleep macrostructure” (Schwarz et al., 2017) might reflect the lack of a reliable measure of sleep cycles.

Revised version: Moreover, a larger difference in classical vs fractal cycle duration was associated with a higher proportion of wake after sleep onset (WASO) in 3/5 datasets as well as in the merged dataset (Table 3). On the other hand, the partial correlation between fractal cycle duration and age remained significant after controlling for the WASO amount. This hints that the association between fractal cycles and age might reflect more than just a confounding effect of WASO. This interpretation is in line with literature on age-related changes in aperiodic activity, namely, on flattering of fractal slopes with age (Voytek et al., 2015; Bódizs et al., 2021; Pathania et al., 2022), especially during SWS (Schneider et al., 2022). Likewise, aging is associated with shorter and fewer classical cycles, with a mean of 3.5 cycles per night compared to the usual 4 – 5 in adults and adolescents (Conte et al., 2014). Our findings suggest that fractal cycles are more sensitive to these age-related alterations than the classical ones. We further speculate that the claim that “age affects sleep microstructure more than sleep macrostructure” (Schwarz et al., 2017) might reflect the lack of a reliable measure of sleep cycles.

Earlier version: Another plausible explanation for longer fractal cycles in younger compared to older adults could be rooted in increased sleep intensity of the younger adults (Jenni & Carskadon, 2004). Further, high sleep intensity driven by homeostatic pressure is associated with the delay in the emergence of the REM sleep phase (Le Bon, 2020; Tarokh et al., 2012). In our dataset, REM latency also decreased with age. Thus, Fig.S5 D (Supplementary Material 2) illustrates that young adults might present with very delayed REM latency, i.e., 200 – 250 minutes after sleep onset, in line with the notion that younger adults more often show cycles with skipped REM sleep (Fig.S6). This can be partly explained by the fact that younger people often have a later chronotype (“night owls”) than older people with puberty linked to delays in the sleep cycle by up to 2 hours (Randler et al., 2016). Young people also have a longer circadian rhythm (> 24 h) than older ones (< 24 h, Monk et al., 2005).

Revised version: Another plausible explanation for longer fractal cycles in younger compared to older adults could be rooted in increased sleep intensity of the younger adults (Jenni & Carskadon, 2004). Further, high sleep intensity driven by homeostatic pressure is associated with the delay in the emergence of the REM sleep phase (Le Bon, 2020; Tarokh et al., 2012). In our dataset, REM latency also decreased with age. Thus, Fig.S5 D (Supplementary Material) illustrates that young adults might present with very delayed REM latency, i.e., 200 – 250 minutes after sleep onset, in line with the notion that younger adults more often show cycles with skipped REM sleep (Fig.S6). This can be partly explained by the fact that younger people often have a later chronotype (“night owls”) than older people with puberty linked to delays in the sleep cycle by up to 2 hours (Randler et al., 2016). Young people also have a longer circadian rhythm (> 24 h) than older ones (< 24 h, Monk et al., 2005).

Earlier version: To further strengthen this line of explanations, we performed a supplemental analysis, which showed that prolonged REM latencies are indeed associated with longer fractal cycles (Fig.S5 C, Supplementary Material 2). Nevertheless, the correlation was weak (yet significant) and observed in the pooled dataset only, i.e., not while analyzing individual datasets. Likewise, the partial correlation between the fractal cycle duration and REM latency adjusted for the participants’ age was non-significant. Moreover, we found that children and adolescents (the group that has the longest REM latencies and the highest rate of cycles with skipped REM sleep) showed shorter fractal cycles compared to young adults, specifically the early-night fractal cycles. In view of these analyses, our attempt to explain longer fractal cycles in younger compared to older adults by increased REM sleep latency becomes less convincing. Moreover, given that our algorithm does not miss cycles with skipped REM sleep, longer REM sleep latencies should not necessarily be related to longer cycles. To summarize, at this stage, the mechanism underlying age-related differences in fractal cycle duration is unclear (possibly with some non-linearities) and future studies are needed to corroborate and further explore it.

Revised version: To further strengthen this line of explanations, we performed a supplemental analysis, which showed that prolonged REM latencies are indeed associated with longer fractal cycles (Fig.S5 C, Supplementary Material). Nevertheless, the correlation was weak (yet significant) and observed in the pooled dataset only, i.e., not while analyzing individual datasets. Likewise, the partial correlation between the fractal cycle duration and REM latency adjusted for the participants’ age was non-significant. Moreover, we found that children and adolescents (the group that has the longest REM latencies and the highest rate of cycles with skipped REM sleep) showed shorter fractal cycles compared to young adults, specifically the early-night fractal cycles. In view of these analyses, our attempt to explain longer fractal cycles in younger compared to older adults by increased REM sleep latency becomes less convincing. Moreover, given that our algorithm does not miss cycles with skipped REM sleep, longer REM sleep latencies should not necessarily be related to longer cycles. To summarize, at this stage, the mechanism underlying age-related differences in fractal cycle duration is unclear (possibly with some non-linearities) and future studies are needed to corroborate and further explore it.

Fractal cycles in MDD

In addition, our study shows that deviations from the observed fractal patterns have some clinical relevance. We found that MDD patients in the medicated state had longer fractal cycles compared to their own unmedicated state and healthy controls. The largest differences were observed for the first sleep cycles. Moreover, patients who took REM-suppressive antidepressants showed prolonged fractal cycles compared to patients who took REM-non-suppressive antidepressants. Given that the fractal cycle duration was defined as an interval of time between two adjacent peaks and that the peaks usually coincide with REM sleep (Fig.2 A), this finding may reflect such aftereffects of antidepressants as delayed onset and reduced amount of REM sleep (Palagini et al., 2013). In other words, if a patient has fewer REM sleep episodes, then the time series of their fractal slopes has fewer peaks and the algorithm detects fewer cycles per total sleep time, i.e., cycle’s duration is longer (Fig.4 A).

Earlier version: Another explanation considers our previous finding that medicated MDD patients show flatter average fractal slopes compared to controls and their own unmedicated state during all sleep stages (Rosenblum et al., 2023 a). This might mean that the antidepressant intake results in shallower fractal fluctuations, which in turn implies that fewer peaks could be detected by our algorithm as the peak threshold was defined a priori in a healthy – not MDD – sample. Interestingly, recently, flatter fractal slopes during REM sleep have been also associated with sustained polyphasic sleep restriction in health (Rosenblum et al., 2024 b), whereas flatter fractal slopes during non-REM sleep were observed in patients with objective insomnia and sleep state misperception, reflecting an abnormally high level of excitation in line with the hyperarousal model of insomnia (Andrillon et al., 2020). Our pilot findings have shown that patient with psychophysiological insomnia have shorter fractal cycles compared to controls (Fig.S10, Supplementary Material 2).

Revised version: Another explanation considers our previous finding that medicated MDD patients show flatter average fractal slopes compared to controls and their own unmedicated state during all sleep stages (Rosenblum et al., 2023a). This might mean that the antidepressant intake results in shallower fractal fluctuations, which in turn implies that fewer peaks could be detected by our algorithm as the peak threshold was defined a priori in a healthy – not MDD – sample. Interestingly, recently, flatter fractal slopes during REM sleep have been also associated with sustained polyphasic sleep restriction in health (Rosenblum et al., 2024b), whereas flatter fractal slopes during non-REM sleep were observed in patients with objective insomnia and sleep state misperception, reflecting an abnormally high level of excitation in line with the hyperarousal model of insomnia (Andrillon et al., 2020). Our pilot findings have shown that patient with psychophysiological insomnia have shorter fractal cycles compared to controls (Fig.S10, Supplementary Material).

Limitations and strengths

The major limitation of this study is its correlational approach, and thus an inability to shed light on the mechanism underlying sleep cycle generation. Therefore, the question of what determines the number and duration of cycles per night remains open. Moreover, further work is needed to determine the mathematically precise and physiologically meaningful model of fractal cycles. Notably, here, we suggest that fractal cycles are a new tool to study the macrostructure of sleep; however, they are presumably not a substitute for the study of the individual sleep stages and microstructural features of sleep (e.g., microarousals, spindles, slow waves).

Additionally, we explored the effect of developmental changes and aging on fractal cycles using a cross-sectional observational approach, whereas these factors might be disentangled more precisely in a longitudinal approach. The age of the pediatric group ranged from 8 – 17 years old; studying younger children and babies would add crucial information on the influence of neurodevelopmental changes on fractal cycles.

The strengths of this study are its large sample size, scripts and data sharing and self-replications in several clinical and healthy datasets of participants in a broad age range, affirming the overall robustness of the phenomena of fractal cycles. Another strength of this work is its generalizability as it has shown that the studies conducted in different experimental environments (including one study conducted at home) using different EEG devices provide comparable results.

To summarize, the large sample and self-replication performed in this study suggest that the “fractal cycle” is a universal concept that should be extensively studied. Displaying the data in the format of fractal cycles provides an intuitive and biologically plausible way to present whole-night sleep neural activity and also adds some graduality to the purely categorical concept of sleep stages that comprise a hypnogram. In future studies, this graduality might help to illuminate differences in sleep architecture across different species, advance our understanding of the role of sleep in neurocognitive development in infants and adolescents as well as in neurodegenerative processes and other fields of neuroscience.

Conclusion

We observed that the slopes of the fractal (aperiodic) spectral power descend and ascend cyclically across a night such that the peaks of the time series of the fractal slopes coincide with REM sleep or sleep lightening while the troughs of these time series coincide with non-REM sleep. Based on this observation, we introduced a new concept of fractal activity-based cycles of sleep or “fractal cycles” for short, defining it as a time interval between two adjacent local peaks of the fractal time series. We have shown that fractal cycles defined by our algorithm largely coincide with classical (non-REM – REM) sleep cycles defined by a hypnogram and replicated our findings in several independently collected healthy and clinical datasets. Moreover, we found that the fractal cycle algorithm reliably detected cycles with skipped REM sleep. In addition, we observed that fractal cycle duration changes as a non-linear function of age, being shorter in children and adolescents compared to young adults as well as in older compared to younger adults. To this end, we conclude that the fractal cycle is an objective, quantifiable and universal concept that could be used to define sleep cycles and display the whole-night sleep neural activity in a more intuitive and biologically plausible way as compared to the conventionally used hypnograms. Having shown that the fractal cycles are prolonged in medicated patients with MDD, we suggest that fractal cycles are a useful tool to study the effects of antidepressants on sleep. Possibly, fractal cycles also will be able to serve as a means to explore sleep architecture alterations in different clinical populations (e.g., to detect REM sleep without atonia) and during neurocognitive development. In summary, this study shows that the fractal cycles of sleep are a promising research tool relevant in health and disease that should be extensively studied.